Population analyses reveal heterogenous encoding in the medial prefrontal cortex during naturalistic foraging

Author response:

The following is the authors’ response to the original reviews

Reviewer 1 (Public Review):

Thank you for the helpful comments. Below, we have quoted the relevant sections from the revised manuscript as we respond to the reviewer’s comments item-by-item.

Weaknesses:

While the task design in this study is intentionally stimulus-rich and places a minimal constraint on the animal to preserve naturalistic behavior, this is, unfortunately, a double-edged sword, as it also introduces additional variables that confound some of the neural analysis. Because of this, a general weakness of the study is a lack of clear interpretability of the task variable neural correlates. This is a limitation of the task, which includes many naturally correlated variables - however, I think with some additional analyses, the authors could strengthen some of their core arguments and significantly improve clarity.

We acknowledge the weakness and have included additional analyses to compensate for it. The details are as follows in our reply to the subsequent comments.

For example, the authors argue, based on an ANN decoding analysis (Figure 2b), that PFC neurons encode spatial information - but the spatial coordinate that they decode (the distance to the active foraging zone) is itself confounded by the fact that animals exhibit different behavior in different sections of the arena. From the way the data are presented, it is difficult to tell whether the decoder performance reflects a true neural correlate of distance, or whether it is driven by behavior-associated activity that is evoked by different behaviors in different parts of the arena. The author's claim that PFC neurons encode spatial information could be substantiated with a more careful analysis of single-neuron responses to supplement the decoder analysis. For example, 1) They could show examples of single neurons that are active at some constant distance away from the foraging site, regardless of animal behavior, and 2) They could quantify how many neurons are significantly spatially modulated, controlling for correlates of behavior events. One possible approach to disambiguate this confound could be to use regression-based models of neuron spiking to quantify variance in neuron activity that is explained by spatial features, behavioral features, or both.

First of all, we would like to point out that while the recording was made during naturalistic foraging with minimal constraints behaviorally, a well-trained rat displayed an almost fixed sequence of actions within each zone. The behavioral repertoire performed in each zone was very different from each other: exploratory behaviors in the N-zone, navigating back and forth in the F-zone, and licking sucrose while avoiding attacks in the E-zone. Therefore, the entire arena is not only divided by the geographical features but also by the distinct set of behaviors performed in each zone. This is evident in the data showing a higher decoding accuracy of spatial distance in the F-zone than in the N- or E-zone. In this sense, the heterogeneous encoding reflects heterogenous distribution of dominant behaviors (navigation in the F-zone and attack avoidance while foraging in the E-zone) and hence corroborate the reviewer’s comment at a macroscopic scale encompassing the entire arena.

Having said that, the more critical question is whether the neural activity is more correlated with microscopic behaviors at every moment rather than the location decoded in the F-zone. As the reviewer suggested, the first-step is to analyze single-neuron activity to identify whether direct neural correlates of location exist. To this end, traditional place maps were constructed for individual neurons. Most neurons did not show cohesive place fields across different regions, indicating little-to-no direct place coding by individual neurons. Only a few neurons displayed recognizable place fields in a consistent manner. However, even these place fields were irregular and patchy, and therefore, nothing comparable to the place cells or grid cells found in the hippocampus or entorhinal cortex. Some examples firing maps have been added to Figure 2 and characterized in the text as below.

“To determine whether location-specific neural activity exists at the single-cell level in our mPFC data, a traditional place map was constructed for individual neurons. Although most neurons did not show cohesive place fields across different regions in the arena, a few neurons modulated their firing rates based on the rat’s current location. However, even these neurons were not comparable to place cells in the hippocampus (O’Keefe & Dostrovsky, 1971) or grid cells in the entorhinal cortex (Hafting et al., 2005) as the place fields were patchy and irregular in some cases (Figure 2B; Units 66 and 125) or too large, spanning the entire zone rather than a discrete location within it (Units 26 and 56). The latter type of neuron has been identified in other studies (e.g., Kaefer et al., 2020).”

Next, to verify whether the location decoding reflects neuronal activity due to external features or particular type of action, predicted location was compared between the opposite directions within the F-zone, inbound and outbound in reference to the goal area (Lobsterbot). If the encoding were specifically tied to a particular action or environmental stimuli, there should be a discrepancy when the ANN decoder trained with outbound trajectory is tested for predictions on the inbound path, and vice versa. However, the results showed no significant difference between the two trajectories, suggesting that the decoded distance was not simply reflecting neural responses to location-specific activities or environmental cues during navigation.

“To determine whether the accuracy of the regressor varied depending on the direction of movement, we compared the decoding accuracy of the regressor for outbound (from the N- to E-zone) vs. inbound (from the E- to N- zone) navigation within the F-zone. There was no significant difference in decoding accuracy between outbound vs. inbound trips (paired t-test; t(39) = 1.52, p =.136), indicating that the stability of spatial encoding was maintained regardless of the moving direction or perceived context (Figure 2E).”

Additionally, we applied the same regression analysis on a subset of data that were recorded while the door to the robot compartment was closed during the Lobsterbot sessions. This way, it is possible to test the decoding accuracy when the most salient spatial feature, the Lobsterbot, is blocked out of sight. The subset represents an average of 38.92% of the entire session. Interestingly, the decoding accuracy with the subset of data was higher accuracy than that with the entire dataset, indicating that the neural activities were not driven by a single salient landmark. This finding supports our conclusion that the location information can be decoded from a population of neurons rather than from individual neurons that are associated with environmental or proprioceptive cues. We have added the following description of results in the manuscript.

“Previous analyses indicated that the distance regressor performed robustly regardless of movement direction, but there is a possibility that the decoder detects visual cues or behaviors specific to the E-zone. For example, neural activity related to Lobsterbot confrontation or licking behavior might be used by the regressor to decode distance. To rule out this possibility, we analyzed a subset of data collected when the compartment door was closed, preventing visual access to the Lobsterbot and sucrose port and limiting active foraging behavior. The regressor trained on this subset still decoded distance with a MAE of 12.14 (± 3.046) cm (paired t-test; t(39) = 12.17, p <.001). Notably, the regressor's performance was significantly higher with this subset than with the full dataset (paired t-test; t(39) = 9.895, p <.001).”

As for the comment on “using regression-based models of neuron spiking to quantify variance in neuron activity that is explained by spatial features, behavioral features, or both”, it is difficult to separate a particular behavioral event let alone timestamping it since the rat’s location was being monitored in the constantly-moving, naturalistic stream of behaviors. However, as mentioned above, a new section entitled “Overlapping populations of mPFC neurons adaptively encode spatial information and defensive decision” argues against single-neuron based account by performing the feature importance analysis. The results showed that even when the top 20% of the most informative neurons were excluded, the remaining neural population could still decode both distance and events. This analysis supports the idea of a population-wide mode shift rather than distinct subgroups of neurons specialized in processing different sensory or motor events. This idea is also expressed in the schematic diagrams featured in Figure 8 of the revision.

To substantiate the claim that PFC neurons really switch between different coding "modes," the authors could include a version of this analysis where they have regressed out, or otherwise controlled for, these confounds. Otherwise, the claim that the authors have identified "distinctively different states of ensemble activity," as opposed to simple coding of salient task features, seems premature.

A key argument in our study is that the mPFC neurons encode different abstract internal representations (distance and avoidance decision) at the level of population. This has been emphasized in the revision with additional analyses and discussions. Most of all, we performed single neuron-based analysis for both spatial encoding (place fields for individual neurons) and avoidance decision (PETHs for head entry and head withdrawal) and contrasted the results with the population analysis. Although some individual neurons displayed a fractured “place cell-like” activity, and some others showed modulated firing at the head-entry and the head-withdrawal events, the ensemble decoding extracted distance information for the current location of the animal at a much higher accuracy. Furthermore, the PCA analysis identified abstract feature dimensions especially regarding the activity in the E-zone that cannot be attributable to a small number of sensory- or motor-related neurons.

To mitigate the possibility that the PCA is driven primarily by a small subset of units responsive to salient behavioral events, we also applied PCA to the dataset excluding the activity in the 2-second time window surrounding the head entry and withdrawal. While this approach does not eliminate all cue- or behavior-related activity within the E-zone, it does remove the neural activity associated with emotionally significant events, such as entry into the E-zone, the first drop of sucrose, head withdrawal, and the attack. Even without these events, the PC identified in the E-zone was still separated from those in the F-zone and N-zone. This result again argues in support of distinct states of ensemble activity formed in accordance with different categories of behaviors performed in different zones. Finally, the Naïve Bayesian classifier trained with ensemble activity in the E-zone was able to predict the success and failure of avoidance that occur a few seconds later, indicating that the same population of neurons are encoding the avoidance decision rather than the location of the animal.

Reviewer 1 (Recommendations):

The authors include an analysis (Figure 4) of population responses using PCA on session-wide data, which they use to support the claim that PFC neurons encode distinctive neural states, particularly between the encounter zone and nesting/foraging zones. However, because the encounter zone contains unique stimulus and task events (sucrose, threat, etc.), and the samples for PCA are drawn from the entire dataset (including during these events), it seems likely that the Euclidean distance measures analyzed in Figure 4b are driven mostly by the neural correlates of these events rather than some more general change in "state" of PFC dynamics. This does not invalidate this analysis but renders it potentially redundant with the single neuron results shown in Figure 5 - and I think the interpretation of this as supporting a state transition in the coding scheme is somewhat misleading. The authors may consider performing a PCA/population vector analysis on the subset of timepoints that do not contain unique behavior events, rather than on session-wide data, or otherwise equalizing samples that correspond to behavioral events in different zones. Observing a difference in PC-projected population vectors drawn from samples that are not contaminated by unique encounter-related events would substantiate the idea that there is a general shift in neural activity that is more related to the change in context or goal state, and less directly to the distinguishing events themselves.

Thank you for the comments. Indeed, this is a recurring theme where the reviewers expressed concerns and doubts about heterogenous encoding of different functional modes. Besides the systematic presentation of the results in the manuscript, from PETH to ANN and to Bayesian classifier, we argue, however, that the activity of the mPFC neurons is better represented by the population rather than loose collection of stimulus- or event-related neurons.

The PCA results that we included as the evidence of distinct functional separation, might reflect activities driven by a small number of event-coding neurons in different zones. As mentioned in the public review, we conducted the same analysis on a subset of data that excluded neural activity potentially influenced by significant events in the E-zone. The critical times are defined as ± 1 second from these events and excluded from the neural data. Despite these exclusions, the results continued to show populational differences between zones, reinforcing the notion that neurons encode abstract behavioral states (decision to avoid or stay) without the sensory- or motor-related activity. Although this analysis does not completely eliminate all possible confounding factors emerging in different external and internal contexts, it provides extra support for the population-level switch occurring in different zones.

In Figure 7, the authors include a schematic that suggests that the number of neurons representing spatial information increases in the foraging zone, and that they overlap substantially with neurons representing behaviors in the encounter zone, such as withdrawal. They show in Figure 3 that location decoding is better in the foraging zone, but I could not find any explicit analysis of single-neuron correlates of spatial information as suggested in the schematic. Is there a formal analysis that lends support to this idea? It would be simple, and informative, to include a quantification of the fraction of spatial- and behavior-modulated neurons in each zone to see if changes in location coding are really driven by "larger" population representations. Also, the authors could quantify the overlap between spatial- and behavior-modulated neurons in the encounter zone to explicitly test whether neurons "switch" their coding scheme.

The Figure 7 (now Figure 8) is now completely revised. The schematic diagram is modified to show spatial and avoidance decision encoding by the overlapping population of mPFC neurons (Figure 8a). Most notably, there are very few neurons that encode location but not the avoidance decision or vice versa. This is indicated by the differently colored units in F-zone vs. E-zone. The model also included units that are “not” engaged in any type of encoding or engaged in only one-type of encoding although they are not the majority.

We have also added a schematic for hypothetical switching mechanisms (Figure 8b) to describe the conceptual scheme for the initiation of encoding-mode switching (sensory-driven vs. arbitrator-driven process)

“Two main hypotheses could explain this switch. A bottom-up hypothesis suggests sensory inputs or upstream signals dictate encoding priorities, while a top-down hypothesis proposes that an internal or external “arbitrator” selects the encoding mode and coordinates the relevant information (Figure 8B). Although the current study is only a first step toward finding the regulatory mechanism behind this switch, our control experiment, where rats reverted to a simple shuttling task, provide evidence that might favor the top-down hypothesis. The absence of the Lobsterbot degraded spatial encoding rather than enhancing it, indicating that simply reducing the task demand is not sufficient to activate one particular type of encoding mode over another. The arbitrator hypothesis asserts that the mPFC neurons are called on to encode heterogenous information when the task demand is high and requires behavioral coordination beyond automatic, stimulus-driven execution. Future studies incorporating multiple simultaneous tasks and carefully controlling contextual variables could help determine whether these functional shifts are governed by top-down processes involving specific neural arbitrators or by bottom-up signals.”

Related to this difference in location coding throughout the environment, the authors suggest in Figure 3a-b that location coding is better in the foraging zone compared to the nest or encounter zones, evidenced by better decoder performance (smaller error) in the foraging zone (Figure 3b). The authors use the same proportion of data from the three zones for setting up training/test sets for cross-validation, but it seems likely that overall, there are substantially more samples from the foraging zone compared to the other two zones, as the animal traverses this section frequently, and whenever it moves from the next into the encounter zone (based on the video). What does the actual heatmap of animal location look like? And, if the data are down-sampled such that each section contributes the same proportion of samples to decoder training, does the error landscape still show better performance in the foraging zone? It is important to disambiguate the effects of uneven sampling from true biological differences in neural activity.

Thank you for the comment. We agree with the concern regarding uneven data size from different sections of the arena. Indeed, as the heatmap below indicates, the rats spent most of their time in two critical locations, one being a transition area between N-and F-zone and the other near the sucrose port. This imbalance needs to be corrected. In fact we have included methodology to correct this biased sampling. In the result section “Non-navigational behavior reduces the accuracy of decoded location” we have the following results.

Heatmap of the animal’s position during one example session.

(Left) Unprocessed occupancy plot. Each dot represents 0.2 seconds. Right) Smoothed occupancy plot using a Gaussian filter (sigma: 10 pixels, filter size: 1001 pixels). The white line indicates a 10 cm length.

“To correct for the unequal distribution of location visits (more visits to the F- than to other zones), the regressor was trained using a subset of the original data, which was equalized for the data size per distance range (see Materials and Methods). Despite the correction, there was a significant main effect of the zone (F(1.16, 45.43) = 119.2, p <.001) and the post hoc results showed that the MAEs in the N-zone (19.52 ± 4.46 cm; t(39) = 10.45; p <.001) and the E-zone (26.13 ± 7.57 cm; t(39) = 11.40; p <.001) had a significantly higher errors when compared to the F-zone (14.10 ± 1.64 cm).”

Also in the method section, we have stated that:

“In the dataset adjusted for uneven location visits, we divided distance values into five equally sized bins. Then, a sub-dataset was created that contains an equal number of data points for each of these bins.”

Why do the authors choose to use a multi-layer neural network (Figure 2b-c) to decode the animal's distance to the encounter zone?(…) The authors may consider also showing an analysis using simple regression, or maybe something like an SVM, in addition to the ANN approach.

We began with a simple linear regression model and progressed to more advanced methods, including SVM and multi-layer neural networks. As shown below, simpler methods could decode distance to some extent, but neural networks and random forest regressors outperformed others (Neural Network: 16.61 cm ± 3.673; Linear Regression: 19.85 cm ± 2.528; Quadratic Regression: 18.68 cm ± 4.674; SVM: 18.88 cm ± 2.676; Random Forest: 13.59 cm ± 3.174).

We chose the neural network model for two main reasons: (1) previous studies demonstrated its superior performance compared to Bayesian regressors commonly used for decoding neural ensembles, and (2) its generalizability and robustness against noisy data. Although the random forest regressor achieved the lowest decoding error, we avoided using it due to its tendency to overfit and its limited generalization to unseen data.

Overall, we expect similar results with other regressors but with different statistical power for decoding accuracy. Instead, we speculate that neural network’s use of multiple nodes contributes to robustness against noise from single-unit recordings and enables the network to capture distributed processing within neural ensembles.

In Figure 6c, the authors show a prediction of withdrawal behavior based on neural activity seconds before the behavior occurs. This is potentially very interesting, as it suggests that something about the state of neural dynamics in PFC is potentially related to the propensity to withdraw, or to the preparation of this behavior. However, another possibility is that the behaves differently, in more subtle ways, while it is anticipating threat and preparing withdrawal behavior - since PFC neurons are correlated with behavior, this could explain decoder performance before the withdrawal behavior occurs. To rule out this possibility, it would be useful to analyze how well, and how early, withdrawal success can be decoded only on the basis of behavioral features from the video, and then to compare this with the time course of the neural decoder. Another approach might be to decode the behavior on the basis of video data as well as neural data, and using a model comparison, measure whether inclusion of neural features significantly increases decoder performance.

We appreciate this important point, as mPFC activity might indeed reflect motor preparation preceding withdrawal behavior. Another reviewer raised a similar concern regarding potential micro-behavioral influences on mPFC activity prior to withdrawal responses. However, our behavioral analysis suggests that highly trained rats engage in sucrose licking which has little variability regardless of the subsequent behavioral decision. To support, 95% of inter-lick intervals were less than 0.25 seconds, which is not enough time to perform any additional behavior during encounters.

To further clarify this, we included additional video showing both avoidance and escape withdrawals at close range. This video was recorded during the development of the behavioral paradigm, though we did not routinely collect this view, as animals consistently exhibited stable licking behavior in the E-zone. As demonstrated in the video, the rat remains highly focused on the lick port with minimal body movement during encounters. Therefore, we believe that the neural ensemble dynamics observed in the mPFC are unlikely to be driven by micro-behavioral changes.

Reviewer 2 (Public Review):

Thank you for the positive comment on our behavior paradigm and constructive suggestions on additional analysis. We came to think that the role of mPFC could be better portrayed as representing and switching between different encoding targets under different contexts, which in part, was more clearly manifested by the naturalistic behavioral paradigm. In the revision we tried to convey this message more explicitly and provide a new perspective for this important aspect of mPFC function.

It is not clear what proportion of each of the ensembles recorded is necessary for decoding distance from the threat, and whether it is these same neurons that directly 'switch' to responding to head entry or withdrawal in the encounter phase within the total population. The PCA gets closest to answering this question by demonstrating that activity during the encounter is different from activity in the nesting or foraging zones, but in principle this could be achieved by neurons or ensembles that did not encode spatial parameters. The population analyses are focused on neurons sensitive to behaviours relating to the threat encounter, but even before dividing into subtypes etc., this is at most half of the recorded population.

In our study, the key idea we aim to convey is that mPFC neurons adapt their encoding schemes based on the context or functional needs of the ongoing task. Other reviewers also suggested strengthening the evidence that the same neurons directly switch between encoding two different tasks. The counteracting hypothesis to "switching functions within the same neurons" posits that there are dedicated subsets of neurons that modulate behavior—either by driving decisions/behaviors themselves or being driven by computations from other brain regions.

To test this idea, we included an additional analysis chapter in the results section titled Overlapping populations of mPFC neurons adaptively encode spatial information and defensive decision. In this section, we directly tested this hypothesis by examining each neuron's contribution to the distance regressor and the event classifier. The results showed that the histogram of feature importance—the contribution to each task—is highly skewed towards zero for both decoders, and removing neurons with high feature importance does not impair the decoder’s performance. These findings suggest that 1) there is no direct division among neurons involved in the two tasks, and 2) information about spatial/defensive behavior is distributed across neurons.

Furthermore, we tested whether there is a negative correlation between the feature importance of spatial encoding and avoidance encoding. Even if there were no “key neurons” that transmit a significant amount of information about either spatial or defensive behavior, it is still possible that neurons with higher information in the navigation context might carry less information in the active-foraging context, or vice versa. However, we did not observe such a trend, suggesting that mPFC neurons do not exhibit a preference for encoding one type of information over the other.

Lastly, another reviewer raised the concern that the PCA results, which we used as evidence of functional separation of different ensemble functions, might be driven by a small number of event-coding neurons. To address this, we conducted the same analysis on a subset of data that excluded neural activity potentially influenced by significant events in the E-zone. In the Peri-Event Time Histogram (PETH) analysis, we observed that some neurons exhibit highly-modulated activity upon arrival at the E-zone (head entry; HE) and immediately following voluntary departure or attack (head withdrawal; HW). We defined 'critical event times' as ± one second from these events and excluded neural data from these periods to determine if PCA could still differentiate neural activities across zones. Despite these exclusions, the results continued to show populational differences between zones, reinforcing the notion that neurons adapt their activity according to the context. We acknowledge that this analysis still cannot eliminate all of the confounding factors due to the context change, but we confirmed that excluding two significant events (delivery onset of sucrose and withdrawal movement) does not alter our result.

To summarize, these additional results further support the conclusion that spatial and avoidance information is distributed across the neural population rather than being handled by distinct subsets. The analyses revealed no negative correlation between spatial and avoidance encoding, and excluding event-driven neural activity did not alter the observed functional separation, confirming that mPFC neurons dynamically adjust their activity to meet contextual demands.

A second concern is also illustrated by Fig. 7: in the data presented, separate reward and threat encoding neurons were not shown - in the current study design, it is not possible to dissociate reward and threat responses as the data without the threat present were only used to study spatial encoding integrity.

Thank you for this valuable feedback. Other reviewers have also noted that Figure 7 (now Figure 8) is misleading and contains assertions not supported by our experiments. In response, we have revised the model to more accurately reflect our findings. We have eliminated the distinction between reward coding and threat coding neurons, simplifying it to focus on spatial encoding and avoidance encoding neurons. The updated figure will more appropriately align with our findings and claims. A. Distinct functional states (spatial vs. avoidance decision) encoded by the same population neurons are separable by the region (F- vs. E zone). B. Hypothetical control models by which mPFC neurons assume different functional states.

Thirdly, the findings of this work are not mechanistic or functional but are purely correlational. For example, it is claimed that analyzing activity around the withdrawal period allows for ascertaining their functional contributions to decisions. But without a direct manipulation of this activity, it is difficult to make such a claim. The authors later discuss whether the elevated response of Type 2 neurons might simply represent fear or anxiety motivation or threat level, or whether they directly contribute to the decision-making process. As is implicit in the discussion, the current study cannot differentiate between these possibilities. However, the language used throughout does not reflect this.

We acknowledge that our experiments only involve correlational study and this serves as weakness. Although we carefully managed to select word to not to be deterministic, we agree that some of the language might mislead readers as if we found direct functional contribution. Thus, we changed expressions as below.

“We then further analyzed the (functional contribution ->)correlation between neural activity and success and failure of avoidance behavior. If the mPFC neurons (encode ->)participate in the avoidance decisions, avoidance withdrawal (AW; withdrawal before the attack) and escape withdrawal (EW; withdrawal after the attack) may be distinguishable from decoded population activity even prior to motor execution.”

Also, we added part below in discussion section to clarify the limitations of the study.

“Despite this interesting conjecture, any analysis based on recording data is only correlational, mandating further studies with direct manipulation of the subpopulation to confirm its functional specificity.”

Fourthly, the authors mention the representation of different functions in 'distinct spatiotemporal regions' but the bulk of the analyses, particularly in terms of response to the threat, do not compare recordings from PL and IL although - as the authors mention in the introduction - there is prior evidence of functional separation between these regions.

Thank you for bringing this part to our attention. As we mentioned in the introduction, we acknowledge the functional differences between the PL and IL regions. Although differences in spatial encoding between these two areas were not deeply explored, we anticipated finding differences in event encoding, given the distinct roles of the PL and IL in fear and threat processing. However, our initial analysis revealed no significant differences in event encoding between the regions, and as a result, we did not emphasize these differences in the manuscript. To address this point, we have reanalyzed the data separately and included the following findings in the manuscript.

“However, we did not observe a difference in decoding accuracy between the PL and IL ensembles, and there were no significant interactions between regressor type (shuffled vs. original) and regions (mixed-effects model; regions: p=.996; interaction: p=.782). These results indicate that the population activity in both the PL and IL contains spatial information (Figure 2D, Video 3).

[…]

Furthermore, we analyzed whether there is a difference in prediction accuracy between sessions with different recorded regions, the PL and the IL. A repeated two-way ANOVA revealed no significant difference between recorded regions, nor any interaction (regions: F(1, 38) = 0.1828, p = 0.671; interaction: F(1, 38) = 0.1614, p = 0.690).

[…]

We also examined whether there is a significant difference between the PL and IL in the proportion of Type 1 and Type 2 neurons. In the PL, among 379 recorded units, 143 units (37.73%) were labeled as Type 1, and 75 units (19.79%) were labeled as Type 2. In contrast, in the IL, 156 units (61.66%) and 19 units (7.51%) of 253 recorded units were labeled as Type 1 and Type 2, respectively. A Chi-square analysis revealed that the PL contains a significantly higher proportion of Type 2 neurons (χ²(1, 632) = 34.85, p < .001), while the IL contains a significantly higher proportion of Type 1 neurons compared to the other region (χ²(1, 632) = 18.07, p < .001).”

To summarize our additional results, we did not observe performance differences in distance decoding or event decoding. The only difference we observed was the proportional variation of Type 1 and Type 2 neurons when we separated the analysis by brain region. These results are somewhat counterintuitive, considering the distinct roles of the two regions—particularly the PL in fear expression and the IL in extinction learning. However, since the studies mentioned in the introduction primarily used lesion and infusion methods, this discrepancy may be due to the different approach taken in this study. Considering this, we have added the following section to the discussion.

“Interestingly, we found no difference between the PL and IL in the decoding accuracy of distance or avoidance decision. This somewhat surprising considering distinct roles of these regions in the long line of fear conditioning and extinction studies, where the PL has been linked to fear expression and the IL to fear extinction learning (Burgos-Robles et al., 2009; Dejean et al., 2016; Kim et al., 2013; Quirk et al., 2006; Sierra-Mercado et al., 2011; Vidal-Gonzalez et al., 2006). On the other hand, more Type 2 neurons were found in the PL and more Type 1 neurons were found in the IL. To recap, typical Type 1 neurons increased the activity briefly after the head entry and then remained inhibited, while Type 2 neurons showed a burst of activity during head entry and sustained increased activity. One study employing context-dependent fear discrimination task (Kim et al., 2013) also identified two distinct types of PL units: short-latency CS-responsive units, which increased firing during the initial 150 ms of tone presentation, and persistently firing units, which maintained firing for up to 30 seconds. Given the temporal dynamics of Type 2 neurons, it is possible that our unsupervised clustering method may have merged the two types of neurons found in Kim et al.’s study.

While we did not observe decreased IL activity during dynamic foraging, prior studies have shown that IL excitability decreases after fear conditioning (Santini et al., 2008), and increased IL activity is necessary for fear extinction learning. In our paradigm, extinction learning was unlikely, as the threat persisted throughout the experiment. Future studies with direct manipulation of these subpopulations, particularly examining head withdrawal timing after such interventions, could provide insight into how these subpopulations guide behavior.”

Additionally, we made some changes in the introduction, mainly replacing the PL/IL with mPFC to be consistent with the main body of results and conclusion and also specifying the correlational nature of the recording study.

“Machine learning-based populational decoding methods, alongside single-cell analyses, were employed to investigate the correlations between neuronal activity and a range of behavioral indices across different sections within the foraging arena.”

Reviewer 2 (Recommendations):

The authors consistently use parametric statistical tests throughout the manuscript. Can they please provide evidence that they have checked whether the data are normally distributed? Otherwise, non-parametric alternatives are more appropriate.

Thank you for mentioning this important issue in the analysis. We re-ran the test of normality for all our data using the Shapiro-Wilk test with a p-value of .05 and found that the following data sets require non-parametric tests, as summarized in Author response table 1 below. For those analyses which did not pass the normality test, we used a non-parametric alternative test instead. We also updated the methods section. For instance, repeated measures ANOVA for supplementary figure S1 and PCA results were changed to the Friedman test with Dunn’s multiple comparison test.

Line 107: it is not clear here or in the methods whether a single drop of sucrose solution is delivered per lick or at some rate during the encounter, both during the habituation or in the final task. This is important information in order to understand how animals might make decisions about whether to stay or leave and how to interpret neural responses during this time period. Or is it a large drop, such that it takes multiple licks to consume? Please clarify.

The apparatus we used incorporated an IR-beam sensor-controlled solenoid valve. As the beam sensor was located right in front of the pipe, the rat’s tongue activated the sensor. As a result, each lick opened the valve for a brief period, releasing a small amount of liquid, and the rat had to continuously lick to gain access to the sucrose. We carefully regulated the flow of the liquid and installed a small sink connected to a vacuum pump, so any remaining sucrose not consumed by the rat was instantly removed from the port. We clarified how sucrose was delivered in the methods section and also in the results section.

Method:

“The sucrose port has an IR sensor which was activated by a single lick. The rat usually stays in front of the lick port and continuously lick up to a rate of 6.3 times per second to obtain sucrose. Any sucrose droplets dropped in the bottom sink were immediately removed by negative pressure so that the rat’s behavior was focused on the licking.”

Result:

“The lick port was activated by an IR-beam sensor, triggering the solenoid valve when the beam was interrupted. The rat gradually learned to obtain rewards by continuously licking the port.”

However, I'm not sure I understand the authors' logic in the interpretation: does the S-phase not also consist of goal-directed behaviour? To me, the core difference is that one is mediated by threat and the other by reward. In addition, it would be helpful to visualize the behaviour in the S-phase, particularly the number of approaches. This difference in the amount of 'experience' so to speak might drive some of the decrease in spatial decoding accuracy, even if travel distance is similar (it is also not clear how travel distance is calculated - is this total distance?) Ideally, this would also be included as a predictor in the GLM.

We agree that the behaviors observed during the shuttling phase can also be considered goal-directed, as the rat moves purposefully toward explicit goals (the sucrose port and the N-zone during the return trip). However, we argue that there is a significant difference in the level of complexity of these goals.

During the L-phase, the rat not only has to successfully navigate to the E-zone for sucrose but also pay attention to the robots, either to avoid an attack from the robot's forehead or escape the fast-striking motion of the claw. When the rat runs toward the E-zone, it typically takes a side-approaching path, similar to Kim and Choi (2018), and exhibits defensive behaviors such as a stretched posture, which were not observed in the S-phase. This behavioral characteristic differs from the S-phase, where the rat adopted a highly stereotyped navigation pattern fairly quickly (within 3 sessions), evidenced by more than 50 shuttling trajectories per session. In this phase, the rat exhibited more stimulus-response behavior, simply repeating the same actions over time without deliberate optimization.

In our additional experiment with two different levels of goal complexity (reward-only vs. reward/threat conflict), we used a between-subject design in which both groups experienced both the S-phase and L-phase before surgery and underwent only one type of session afterward. This approach ruled out the possibility of differences in contextual experience. Additionally, since we initially designed the S-phase as extended training, behaviors in the apparatus tended to stabilize after rats completed both the S-phase and L-phase before surgery. As a result, we compared the post-surgery Lobsterbot phase to the post-surgery shuttling phase to investigate how different levels of goal complexity shape spatial encoding strength.

To clarify our claim, we edited the paragraph below.

“This absence of spatial correlates may result from a lack of complex goal-oriented navigation behavior, which requires deliberate planning to acquire more rewards and avoid potential threats.

[…]

After the surgery, unlike the Lob-Exp group, the Ctrl-Exp group returned to the shuttling phase, during which the Lobsterbot was removed. With this protocol, both groups experienced sessions with the Lobsterbot, but the Ctrl-Exp group's task became less complex, as it was reduced to mere reward collection.

. Given these observations, along with the mPFC’s lack of consistency in spatial encoding, it is plausible that the mPFC operates in multiple functional modes, and the spatial encoding mode is preempted when the complexity of the task requires deliberate spatial navigation.”

Additionally, we added behavior data during initial S-phase into Supplementary Figure 1.

It is good point that the amount of experience might drive decrease in spatial decoding accuracy. To test this hypothesis, we added a new variable, the number of Lobsterbot sessions after surgery, to the previous GLM analysis. The updated model predicted the outcome variable with significant accuracy (F(4,44) = 10.31, p < .001), and with the R-squared value at 0.4838. The regression coefficients were as follows: presence of the Lobsterbot (2.76, standard error [SE] = 1.11, t = 2.42, p = .020), number of recorded cells (-0.43, SE = .08, t = -5.22, p < .001), recording location (0.90, SE = 1.11, p = .424), and number of L sessions (0.002, SE = 0.11, p = .981). These results indicate that the number of exposures to the Lobsterbot sessions, as a measure of experience, did not affect spatial decoding accuracy.

For minor edit, we edited the term as “total travel distance”.

Relating to the previous point, it should be emphasized in both sections on removing the Lobsterbot and on non-navigational behaviours that the spatial decoding is all in reference to distance from the threat (or reward location). The language in these sections differs from the previous section where 'distance from the goal' is mentioned. If the authors wish to discuss spatial decoding per se, it would be helpful to perform the same analysis but relative to the animals' own location which might have equal accuracy across locations in the arena. Otherwise, it is worth altering the language in e.g. line 258 onwards to state the fact that distance to the goal is only decodable when animals are actively engaged in the task.

Thank you for this comment, we changed the term as “distance from the conflict zone” or “distance of the rat to the center of the E-zone” to clarify our experiment setup.

In Fig. 5, why is the number of neurons shown in the PETHs less than the numbers shown in the pie charts?

The difference in the number of neurons between the PETHs and the pie charts in Figure 5 is because PETHs are drawn only for 'event-responsive' units. For visualizing the neurons, we selectively included those that met certain criteria described in Method section (Behavior-responsive unit analysis). We have updated the caption for Figure 5 as follows to minimize confusion.

“Multiple subpopulations in the mPFC react differently to head entry and head withdrawal.

(A) Top: The PETH of head entry-responsive units is color-coded based on the Z-score of activity.

(C) The PETH of head withdrawal-responsive units is color-coded based on the Z-score of activity.”

I appreciate the amount of relatively unprocessed data plotted in Figure 5, but it would be great to visualize something similar for AW vs. EW responses within the HW2 population. In other words, what is there that's discernably different within these responses that results in the findings of Fig. 6?

To visualize the difference in neural activity between AW and EW, we included an additional supplementary figure (Supplementary Figure 5). We divided the neurons into Type 1 and Type 2 and plotted PETH during Avoidance Withdrawal (AW) and Escape Withdrawal (EW). Consistent with the results shown in Figure 6d, we could visually observe increased activity in Type 2 neurons before the execution of AW compared to EW. However, we couldn’t find a similar pattern in Type 1 neurons.

On a related note, it would add explanatory power if the authors were able to more tightly link the prediction accuracy of the ensemble (particularly the Type 2 neurons) to the timing of the behaviour. Earlier in the manuscript it would be helpful to show latency to withdraw in AW trials; are animals leaving many seconds before the attack happens, or are they just about anticipating the timing of the attack? And therefore when using ensemble activity to predict the success of the AW, is the degree to which this can be done in advance (as the authors say, up to 6 seconds before withdrawal) also related to how long the animal has been engaged with the threat?

We agree that the timing of head withdrawal, particularly in AW trials, is a critical factor in describing the rat's strategy toward the task. To test whether the rat uses a precise timing strategy—for instance, leaving several seconds before the attack or exploiting the discrete 3- and 6-second attack durations—we plotted all head withdrawal timepoints during the 6-second trials. The distribution was more even, without distinguishable peaks (e.g., at the very initial period or at the 3- or 6-second mark). This indicates a lack of precise temporal strategy by the rat. We included additional data in the supplementary figure (Supplementary Figure 6) and added the following to the results section.

“We monitored all head withdrawal timepoints to assess whether rats developed a temporal strategy to differentiate between the 3-second and 6-second attacks. We found no evidence of such a strategy, as the timings of premature head withdrawals during the 6-second attack trials were evenly distributed (see Supplementary Figure S1).”

As depicted in the new supplementary figure, head withdrawal times during avoidance behavior vary from sub-seconds to the 3- or 6-second attack timepoints. After receiving the reviewer’s comment, we became curious whether there is a decoding accuracy difference depending on how long the animal engaged with the threat. We selected all 6-second attack and avoidance withdrawal trials and checked if correctly classified trials (AW trials classified as AW) had different head withdrawal times—perhaps shorter durations—compared to misclassified trials (AW trials classified as EW). As shown in Author response image 3 below, there was no significant difference between these two types, indicating that the latency of head withdrawal does not affect prediction accuracy.

Finally, there remain some open questions. One is how much encoding strength - of either space or the decision to leave during the encounter - relates to individual differences in animal performance or behaviour, particularly because this seems so variable at baseline. A second is how stable this encoding is. The authors mention that the distance encoding must be stable to an extent for their regressor to work; I am curious whether this stability is also found during the encounter coding, and also whether it is stable across experience. For example, in a session when an individual has a high proportion of anticipatory withdrawals, is the proportion of Type 2 neurons higher?

Thank you for these questions. To recap the number of animals that we used, we used five rats during Lobsterbot experiments, and three rats for control experiment that we removed Lobsterbot after training. Indeed, there were individual differences in performance (i.e. avoidance success rate), number of recorded units (related to the recording quality), and baseline behaviors. To clarify these differences, see author response image 4 below.

We used a GLM to measure how much of the decoder’s accuracy was explained by individual differences. The result showed that 38.96% of distance regressor’s performance, and 12.14% of the event classifier’s performance was explained by the individual difference. Since recording quality was highly dependent on the animals, the high subject variability detected in the distance regression might be attributed to the number of recorded cells. Rat00 which had the lowest average mean absolute error had the highest number of recorded cells at average of 18. Compared to the distance regression, there was less subject variability in event classification. Indeed, the GLM results showed that the variability explained by the number of cells was only 0.62% in event classification.

The reason we mentioned that "distance encoding must be stable for our regressor to work" is entirely based on the population-level analysis. Because we used neural data and behaviors from entire trials within a session, the regressor or classifier would have low accuracy if encoding dynamics changed within the session. In other words, if the way neurons encode avoidance/escape predictive patterns changed within a training set, the classifier would fail to generate an optimized separation function that works well across all datasets.

To further investigate whether changes in experience affect event classification results over time, we plotted an additional graph below. Although there are individual and daily fluctuations in decoding accuracy, there was no observable trend throughout the experiments.

Regarding the correlation between the ratio of avoidance withdrawal and the proportion of Type 2 neurons, we were also curious and analyzed the data. Across 40 sessions, the correlation was -0.0716. For Type 1 neurons, it was slightly higher at 0.1459. We believe this indicates no significant relationship between the two variables.

Minor points:

I struggled with the overuse of acronyms in the paper. Some might be helpful but F-zone/N-zone, for example, or HE/HW, AW/EW are a bit of a struggle. After reading the paper a few times I learned them but a naive reader might need to often refer back to when they were first defined (as I frequently had to).

To increase readability, we removed acronyms that are not often used and changed HE/HW to head-entry/head-withdrawal.

I have a few questions about Figure 1F: in the text (line 150) it says that 'surgery was performed after three L sessions when the rats displayed a range of 30% to 60% AW'. This doesn't seem consistent with what is plotted, which shows greater variability in the proportion of AW behaviours both before and after surgery. It also appears that several rats only experienced two days of the L1 phase; please make clear if so. And finally, what is the line at 50% indicating? Neither the text nor the legend discuss any sort of thresholding at 50%. Instead, it would be best to make the distinction between pre- and post-surgery behaviour visually clearer.

Thank you for pointing out this issue. We acknowledge there was an error in the text description. As noted in the Methods section, we proceeded with surgery after three Lobsterbot sessions. We have removed the incorrect part from the Results section and revised the Methods section for clarity.

“After three days of Lobsterbot sessions, the rats underwent microdrive implant surgery, and recording data were collected from subsequent sessions, either Lobsterbot or shuttling sessions, depending on the experiment. For all post-surgery sessions, those with fewer than 20 approaches in 30 minutes were excluded from further analysis.”

Among the five rats, Rat2 and Rat3 did not approach the robot during the entire Lob2 session, which is why these two rats do not have Lob2 data points. We updated the caption for regarding issue.

Initially, we added a 50% reference line, but we agree it is unnecessary as we do not discuss this reference. We have updated the figure to include the surgery point, as shown in Supplementary Figure 1.

Fig. 2C: each dot is an ensemble of simultaneously recorded neurons, i.e. a subset of the total 800-odd units if I understand correctly. How many ensembles does each rat contribute? Similarly, is this evenly distributed across PL and IL?

Yes, each dot represents a single session, with a total of 40 sessions. Five rats contributed 11, 9, 8, 7, and 5 sessions, respectively. Although each rat initially had more than 10 sessions, we discarded some sessions with a low unit count (fewer than 10 sessions; as detailed in Materials and Methods - Data Collection). We collected 25 sessions from the PL and 15 sessions from the IL. Our goal was to collect more than 200 units per each region.

Please show individual data points for Fig. 2D.

We update the figure with individual data points.

Is there a reason why the section on removing the Lobsterbot (lines 200 - 215) does not have associated MAE plots? Particularly the critical comparison between Lob-Exp and Ctl-Exp.

We intentionally removed some graphs to create a more compact figure, but we appreciate your suggestion and have included the graph in Figure 2.

Some references to supplementary materials are not working, e.g. line 333.

Our submitted version of manuscript had reference error. For the current version, we used plane text, and the references are fixed.

The legend for Supp. Fig. 2B is incorrect.

We greatly appreciate this point. We changed the caption to match the figure.

Reviewer 3 (Public Review):

Thank you for recognizing our efforts in designing an ethologically relevant foraging task to uncover the multiple roles of the mPFC. While we acknowledge certain limitations in our methodology—particularly that we only observed correlations between neural activity and behavior without direct manipulation—we have conducted additional analyses to further strengthen our findings.

Weakness:

The primary concern with this study is the absence of direct evidence regarding the role of the mPFC in the foraging behavior of the rats. The ability to predict heterogeneous variables from the population activity of a specific brain area does not necessarily imply that this brain area is computing or using this information. In light of recent reports revealing the distributed nature of neural coding, conducting direct causal experiments would be essential to draw conclusions about the role of the mPFC in spatial encoding and/or threat evaluation. Alternatively, a comparison with the activity from a different brain region could provide valuable insights (or at the very least, a comparison between PL and IL within the mPFC).

Thank you for the comment. Indeed, the fundamental limitation of the recording study is that it is only correlational, and any causal relationship between neural activity and behavioral indices is only speculative. We made it clearer in the revision and refrained from expressing any speculative ideas suggesting causality throughout the revision. While we did not provide direct evidence that the mPFC is computing or utilizing spatial/foraging information, we based our assertion on previous studies that have directly demonstrated the mPFC's role in complex decision-making tasks (Martin-Fernandez et al., 2023; Orsini et al., 2018; Zeeb et al., 2015) and in certain types of spatial tasks (De Bruin et al., 1994; Sapiurka et al., 2016) . We would like to emphasize that, to the best of our knowledge, there was no previous study which investigated the mPFC function while animal is solving multiple heterogenous problems in semi-naturalistic environment. Therefore, although our recording study only provides speculative causal inference, it certainly provides a foundation for investigating the mPFC function. Future study employing more sophisticated, cell-type specific manipulations would confirm the hypotheses from the current study.

One of the key questions of this manuscript is how multiple pieces of information are represented in the recorded population of neurons. Most of the studies mentioned above use highly structured experimental designs, which allow researchers to study only one function of the mPFC. In the current study, the semi-naturalistic environment allows rats to freely switch between multiple behavioral sets, and our decoding analysis quantitatively assesses the extent to which spatial/foraging information is embedded during these sets. Our goal is to demonstrate that two different task hyperspaces are co-expressed in the same region and that the degree of this expression varies according to the rat’s current behavior (See Figure 8(b) in the revised manuscript).

Alternatively, we added multiple analyses. First, we included a single unit-level analysis looking at the place cell-like property to contrast with the ensemble decoding. Most neurons did not show well-defined place fields although there were some indications for place cell-like property. For example, some neurons displayed fragmented place fields or unusually large place fields only at particular spots in the arena (mostly around the gates). The accuracy from this place information at the single-neuron level is much lower than that acquired from population decoding. Likewise, although there were neurons with modulated firing around the time of particular behavior (head entry and withdrawal), overall prediction accuracy of avoidance decision was much higher when the ensemble-based classifier was applied.

Moreover, given that high-dimensional movement has been shown to be reflected in the neural activity across the entire dorsal cortex, more thorough comparisons between the neural encoding of task variables and movement would help rule out the possibility that the heterogeneous encoding observed in the mPFC is merely a reflection of the rats' movements in different behavioral modes.

Thanks for the comment. We acknowledge that the neural activity may reflect various movement components across different zones in the arena. We performed several analyses to test this idea. First, we want to recap our run-and-stop event analysis may provide an insight regarding whether the mPFC neurons are encoding locations despite the significant motor events. The rats typically move across the F-zone fairly routinely and swiftly (as if they are “running”) to reach the E-zone at which they reduce the moving speed to almost a halt (“stopping”). The PETHs around these critical motor events, however, did not show any significant modulation of neural activity indicating that most neurons we recorded from mPFC did not respond to movement.

We added this analysis to demonstrate that these sudden stops did not evoke the characteristic activation of Type 1 and Type 2 neurons observed during head entry into the E-zone. When we isolated these sudden stops outside the E-zone, we did not observe this neural signature (Supplementary Figure 2).

Second, our PCA results showed that population activity in the E-zone during dynamic foraging behavior was distinct from the activity observed in the N- and F-zones during navigation. However, there is a possibility that the two behaviorally significant events—entry into the E-zone and voluntary or sudden exit—might be driving the differences observed in the PCA results. To account for this, we designated ±1 second from head entry and head withdrawal as "critical event times," excluded the corresponding neural data, and reanalyzed the data. This method removed neural activity associated with sudden movements in specific zones. Despite this exclusion, the PCA still revealed distinct population activity in the E-zone, different from the other zones (Supplementary Figure 4). This result reduces the likelihood that the observed heterogeneous neural activity is merely a reflection of zone-specific movements.

Lastly, the main claim of the paper is that the mPFC population switches between different functional modes depending on the context. However, no dynamic analysis or switching model has been employed to directly support this hypothesis.

Thank you for this comment. Since we did not conduct a manipulation experiment, there is a clear limitation in uncovering how switching occurs between the two task contexts. To make the most of our population recording data, we added an additional results section that examines how individual neurons contribute to both the distance regressor and the event classifier. Our findings support the idea that distance and dynamic foraging information are distributed across neurons, with no distinct subpopulations dedicated to each context. This suggests that mPFC neurons adjust their coding schemes based on the current task context, aligning with Duncan’s (2001) adaptive coding model, which posits that mPFC neurons adapt their coding to meet the task's current demands.

Reviewer 3 (Recommendations):

The evidence for spatial encoding is relatively weak. In the F-zone (50 x 48 cm), the average error was approximately 17 cm, constituting about a third of the box's width and likely not significantly smaller than the size of a rat's body. The errors in the shuffled data are also not substantially greater than those in the original data. An essential test indicates that spatial decoding accuracy decreases when the Losterbot is removed. However, assessing the validity of the results is difficult in the current state. There is no figure illustrating the results, and no statistics are provided regarding the test for matching the number of neurons.

We acknowledge that the average error (~ 17 cm ) measured in our study is relatively large, even though the error is significantly smaller than that by the shuffled control model (22.6 cm). Previous studies reported smaller prediction errors but in different experimental conditions: 16 cm in Kaefer et al. (2020) and less than 10 cm in Ma et al. (2023) and Mashhoori et al. (2018). Most notably, the average number of units used in our study (15.8 units per session) is significantly smaller compared to the previous works, which used 63, 49, and 40 units, respectively. As our GLM results demonstrated, the number of recorded cells significantly influenced decoding accuracy (β = -0.43 cm/neuron). With a similar number of recorded cells, we would have achieved comparable decoding accuracy. In addition, unlike other studies that have employed a dedicated maze such as the virtual track or the 8-shaped maze, we exposed rats to a semi-naturalistic environment where they exhibited a variety of behaviors beyond simple navigation. As argued throughout the manuscript, we believe that the spatial information represented in the mPFC is susceptible to disruption when the animal engages in other activities. A similar phenomenon was reported by Mashhoori et al. (2018), where the decoder, which typically showed a median error of less than 10 cm, exhibited a much higher error—nearly 100 cm—near the feeder location.

As for the reviewer’s request for comparing spatial decoding without the Lobsterbot, we added a new figure to illustrate the spatial decoding results, including statistical details. We also applied a Generalized Linear Model to regress out the effect of the number of recorded neurons and statistically assess the impact of Lobsterbot removal. This adjustment directly addresses the reviewer's request for a clearer presentation of the results and helps contextualize the decoding performance in relation to the number of recorded neurons.

As indicated in the public review, drawing conclusions about the role of the mPFC in navigation and avoidance behavior during the foraging task is challenging due to the exclusively correlational nature of the results. The accuracy in AW/EW discrimination increases a few seconds before the response, implying that changes in mPFC activity precede the avoidance/escape response. However, one must question whether this truly reflects the case. Could this phenomenon be attributed to rats modifying their "micro-behavior" (as evidenced by changes in movement observed in the video) before executing the escape response, and subsequently influencing mPFC activity?

We appreciate the reviewer's thoughtful observation regarding the correlational nature of our results and the potential influence of pre-escape micro-behaviors on mPFC activity. We acknowledge that the increased accuracy in AW/EW discrimination preceding the response could also be correlated with micro-behaviors. However, there is very little room for extraneous behavior other than licking the sucrose delivery port within the E-zone, as the rats are highly trained to perform this stereotypical behavior. To support this, we measured the time delays between licking events (inter-lick intervals). The results show a sharp distribution, with 95% of the intervals falling within a quarter second, indicating that the rats were stable in the E-zone, consistently licking without altering their posture.

To complement the data presented in Author response image 2, a video clip showing a rat engaged in licking behavior was included. We carefully designed the robot compartment and adjusted the distance between the Lobsterbot and the sucrose port to ensure that rats could exhibit only limited behaviors inside the E-zone. The video confirms that no significant micro-behaviors were observed during the rat’s activity in the E-zone.

If mPFC activity indeed switches mode, the results do not clearly indicate whether individual cells are specifically dedicated to spatial representation and avoidance or if they adapt their function based on the current goal. Figure 7, presented as a schematic illustration, suggests the latter option. However, the proportion of cells in the HE and HW categories that also encode spatial location has not been demonstrated. It has also not been shown how the switch is manifested at the level of the population.

Thank you for this comment. As the reviewer pointed out, we suggest that mPFC neurons do not diverge based on their functions, but rather adapt their roles according to the current goal. To support this assertion, we added an additional results section that calculates the feature importance of decoders. This analysis allows us to quantitatively measure each neuron’s contribution to both the distance regressor and the event decoder. Our results indicate that distance and defensive behavior are not encoded by a small subset of neurons; instead, the information is distributed across the population. Shuffling the neural data of a single neuron resulted in a median increase in decoding error of 0.73 cm for the distance regressor and 0.01% for the event decoder, demonstrating that the decoders do not rely on a specific subset of neurons that exclusively encode spatial and/or defensive behavior

Although we found supporting evidence that mPFC neurons encode two different types of information depending on the current context, we acknowledge that we could not go further in answering how this switch is manifested. One simple explanation is that the function is driven by current contextual information and goals—in other words, a bottom-up mechanism. However, in our control experiment, simplifying the navigation task worsened the encoding of spatial information in the mPFC. Therefore, we speculate that an external or internal arbitrator circuit determines what information to encode. A precise temporal analysis of the timepoint when the switch occurs in more controlled experiments might answer these questions. We have added this discussion to the discussion section.

PL and IL are two distinct regions; however, there is no comparison between the two areas regarding their functional properties or the representations of the cells. Are the proportions of cell categories (HE vs HW or HE1 vs HE2, spatial encoding vs no spatial encoding) different in IL and PL? Are areas differentially active during the different behaviors?

Thank you for bringing up this issue. As mentioned in our response to the public review, we included a comparison between the PL and IL regions. While we did not observe any differences in spatial encoding (feature importance scores), the only distinction was in the proportion of Type 1 and Type 2 neurons, as the reviewer suggested. We have incorporated our interpretation of these results into the discussion section.

The results and interpretations of the cluster analysis appear to be highly dependent on the parameters used to define a cluster. For example, the HE2 category includes cells with activity that precedes events and gradually decreases afterward, as well as cells with activity that only follows the events.

We strongly agree that dependency on hyperparameters is a crucial point when using unsupervised clustering methods. To eliminate any subjective criteria in defining clusters, we carefully selected our clustering approach, which requires only two hyperparameters: the number of initial clusters (set to 8) and the minimum number of cells required to be considered a valid cluster (cutoff limit, 50). The rationale behind these choices was: 1) a higher number of initial clusters would fail to generalize neural activity, 2) clusters with fewer than 50 neurons would be difficult to analyze, and 3) to prevent the separation of clusters that show noisy responses to the event.

Author response table 2 shows the differences in the number of cell clusters when we varied these two parameters. As demonstrated, changing these two variables does result in different numbers of clusters. However, when we plotted each cluster type’s activity around head entry (HE) and head withdrawal (HW), an increased number of clusters resulted in the addition of small subsets with low variation in activity around the event, without affecting the general activity patterns of the major clusters.

The example mentioned by the reviewer—possible separation of HE2—appears when using a hyperparameter set those results in 4 clusters, not 3. In this result, 83 units, which were labeled as HE2 in the 3-cluster hyperparameter set, form a new group, HE3 (Group 3). This group of units shows increased activity after head entry and exhibited characteristics similar to HE2, with most of the units classified as HW2, maintaining high activity until head withdrawal. Among the 83 HE3 units, 36 were further classified as HW2, 44 as non-significant, and 3 as HW1. Therefore, we believe this does not affect our analysis, as we observed the separation of two major groups, Type 1 (HE1-HW1) and Type 2 (HE2-HW2), and focused our analysis on these groups afterward.

Despite this validation, there remains a strong possibility that our method might not fully capture small yet significant subpopulations of mPFC units. As a result, we have included a sentence in the methods section addressing the rationale and stability of our approach.

“(Materials and Methods) To compensate for the limited number of neurons recorded per session, the hyperparameter set was chosen to generalize their activity and categorize them into major types, allowing us to focus on neurons that appeared across multiple recording sessions. Although changes in the hyperparameter sets resulted in different numbers of clusters, the major activity types remained consistent (Supplementary Figure S8). However, there is a chance that this method may not differentiate smaller subsets of neurons, particularly those with fewer than 50 recorded neurons.”

Minor points:

Line 333: Error! Reference source not found. This was probably the place for citing Figure S2?

Lines 339, 343: Error! Reference source not found.

Thank you for mentioning these comments. In the new version, all reference functions from Word have been replaced with plain text.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Foraging in the wild requires coordinated switching of critical functions, including goal-oriented navigation and context-appropriate action selection. Nevertheless, few studies have examined how different functions are represented in the brain during naturalistic foraging. To address this question, we recorded multiple single-unit activities from the medial prefrontal cortex (mPFC) of rats seeking a sucrose reward in the presence of a robotic predator (Lobsterbot) that posed periodic threats. Simultaneously recorded ensemble activities from 10-24 neurons were analyzed in reference to various behavioral indices as the animal moved freely between the nest (N) and the goal (E) across the foraging (F) area. In the E-zone, the rat initially received and gradually learned to avoid unpredictable attacks by the Lobsterbot. An artificial neural network, trained with simultaneously recorded neural activity, estimated the rat’s current distance from the Lobsterbot. The accuracy of distance estimation was the highest in the middle F-zone in which the dominant behavior was active navigation. The spatial encoding persisted in the N-zone when non-navigational behaviors such as grooming, rearing, and sniffing were excluded. In contrast, the accuracy decreased as the animal approached the E-zone, when the activity of the same neuronal ensembles was more correlated with dynamic decision-making between food procurement and Lobsterbot evasion. A population-wide analysis confirmed a highly heterogeneous encoding by the region. To further assess the decision-related activity in the E-zone, a naïve Bayesian classifier was trained to predict the success and failure of avoidance behavior. The classifier predicted the avoidance outcome as much as 6 s before the head withdrawal. In addition, two sub-populations of recorded units with distinct temporal dynamics contributed differently to the prediction. These findings suggest that the mPFC neurons may adopt at least two modes of heterogenous encoding that reflect the processing of relevant spatial context and the imminent situational challenge.

Revised version: Foraging in the wild requires coordinated switching of critical functions, including goal-oriented navigation and context-appropriate action selection. Nevertheless, few studies have examined how different functions are represented in the brain during naturalistic foraging. To address this question, we recorded multiple single-unit activities from the medial prefrontal cortex (mPFC) of rats seeking a sucrose reward in the presence of an unpredictable attack posed by a robotic predator (Lobsterbot). Simultaneously recorded ensemble activities from neurons were analyzed in reference to various behavioral indices as the animal moved freely across the foraging area (F) between the nest (N) and the goal (E) area. An artificial neural network, trained with simultaneously recorded neural activity, estimated the rat’s current distance from the Lobsterbot. The accuracy of distance estimation was the highest in the middle F-zone in which the dominant behavior was active navigation. The spatial encoding persisted in the N-zone when non-navigational behaviors such as grooming, rearing, and sniffing were excluded. In contrast, the accuracy decreased as the animal approached the E-zone, when the activity of the same neuronal ensembles was more correlated with events related to dynamic decision-making between food procurement and Lobsterbot evasion. A population-wide analysis confirmed highly heterogeneous encoding by the region. To further assess the decision-related activity in the E-zone, a naïve Bayesian classifier was trained to predict the success and failure of avoidance behavior. The classifier predicted the avoidance outcome as much as 6 s before the head withdrawal. In addition, two sub-populations of recorded units with distinct temporal dynamics contributed differently to the prediction. These findings suggest that an overlapping population of mPFC neurons may switch between two heterogenous modes, encoding relevant locations for goal-directed navigation or an imminent situational challenge.

Introduction

Foraging in nature entails context-appropriate action selection to resolve a mixture of problems such as risk and value assessment, maintenance of goal-oriented navigation, and task-switching (Pyke et al., 1977; Stephens & Krebs, 1986). For example, foraging animals must update the value of the goal in a space-time domain, negotiate a threat while attaining food, and regulate a variety of stimulus-driven defensive behaviors (Fanselow & Lester, 1988; Kim et al., 2018; Mobbs et al., 2007, 2020; Mobbs & Kim, 2015). All these survival-related problems require executive functions that can flexibly update risk and value representations in working memory to adapt to the foraging environment. Numerous studies have identified the prefrontal cortex as the primary computational circuit for hosting multiple strategies as well as relevant survival-related information (Fuster, 1973; Korn & Bach, 2019; Lee et al., 2014).

Earlier version: A critical component of successful foraging is goal-directed navigation. Although the hippocampus (HIPP) is crucial for maintaining reward location and path planning (Jarzebowski et al., 2022; Krishnan et al., 2022; Nyberg et al., 2022; Wikenheiser et al., 2013; Wikenheiser & Redish, 2015), the medial prefrontal cortex (mPFC) and its interaction with the HIPP also plays a role in goal-directed navigation (Guise & Shapiro, 2017; Sapiurka et al., 2016; Spellman et al., 2015; Yu et al., 2018). Furthermore, given that the mPFC is involved in deliberately strategizing defensive behaviors across different situations (see Mobbs et al., 2020 for review), navigation during naturalistic foraging where risk and reward coexist may serve as a suitable context for exploring the flexible information processing capabilities of the mPFC. In fact, the mPFC houses spatially selective neurons that represent relative maze positions (Kaefer et al., 2020; Yang & Mailman, 2018), goal locations (Hok et al., 2005; however see Böhm & Lee, 2020 for contrary results), and precise Cartesian coordinates during food procurement (Mashhoori et al., 2018).

Revised version: A critical component of successful foraging is goal-directed navigation. Although the hippocampus (HIPP) is crucial for maintaining reward location and path planning (Jarzebowski et al., 2022; Krishnan et al., 2022; Nyberg et al., 2022; Wikenheiser et al., 2013; Wikenheiser & Redish, 2015), the medial prefrontal cortex (mPFC) and its interaction with the HIPP also play a role in goal-directed navigation (Guise & Shapiro, 2017; Sapiurka et al., 2016; Spellman et al., 2015; Yu et al., 2018). Furthermore, given that the mPFC is involved in processing deliberate defensive strategies in multiple situations (see Mobbs et al., 2020 for review), naturalistic foraging where risk and reward coexist may serve as a suitable foundation for exploring adaptive information processing in the mPFC. Notably, mPFC houses spatially selective neurons that represent relative positions in a maze (Kaefer et al., 2020; Yang & Mailman, 2018), goal locations (Hok et al., 2005; but see Böhm & Lee, 2020), and precise Cartesian coordinates during food procurement (Mashhoori et al., 2018).

Earlier version: The mPFC plays a role in processing emotional stimuli, especially those signaling threats in the environment. Not only is the mPFC intricately intertwined with the major limbic areas including the amygdala (Gabbott et al., 2005; Hoover & Vertes, 2007; Vertes, 2004), but it is also causally linked to defensive behaviors. For example, the infralimbic cortex (IL) of the mPFC is associated with the extinction of conditioned fear memory (Sierra-Mercado et al., 2006, 2011). In contrast, the activation of the prelimbic cortex (PL) is linked to the expression of fear-conditioned responses (Burgos-Robles et al., 2009; Dejean et al., 2016; M. M. Diehl et al., 2018; Kim et al., 2013). These findings suggest that the critical information for risky decision-making might be represented and processed in the mPFC. Furthermore, PL neurons show increased activity in response to survival-related stimuli such as food, artificial predators, or both (Kim et al., 2018).

Revised version: The mPFC also plays a role in processing emotional stimuli, especially those signaling threats in the environment. In line with its intricate connections with the major limbic areas including the amygdala (Gabbott et al., 2005; Hoover & Vertes, 2007; Vertes, 2004), the mPFC is causally linked to learning and expression of fear responses. For example, the infralimbic cortex (IL) of the mPFC is associated with the extinction of conditioned fear response (Park & Choi, 2010; Sierra-Mercado et al., 2006). In contrast, the activation of the prelimbic cortex (PL) is linked to the expression of fear-conditioned responses (Burgos-Robles et al., 2009; Dejean et al., 2016; M. M. Diehl et al., 2018; Kim et al., 2013). In addition, PL neurons show increased activity in response to survival-related stimuli such as food, artificial predators, or both (Kim et al., 2018). These findings suggest that the critical information for risky decision-making might be represented and processed in the mPFC.

Added in the revised version: Currently, there is limited understanding of how the brain organizes navigation and behavior selections during naturalistic foraging. It is unclear whether spatial representation and emotional processing are separated at the single neuron level or coexist in overlapping populations of the mPFC neurons. In this study, we devised a naturalistic foraging task which allowed self-initiated shuttling between different locations and require avoiding an unpredictable attack by a robotic predator while approaching sucrose. This approach allowed continuous monitoring of the neural activity as the rat switched between navigation and strategic reward procurement. Machine learning-based populational decoding methods, alongside single-cell analyses, were employed to investigate the correlations between neuronal activity and a range of behavioral indices across different sections within the foraging arena.

Removed in the revised version: Currently, there is limited understanding of how information is represented in the mPFC during dynamic foraging. It is unclear whether individual cells are dedicated to spatial representation and emotional stimuli processing or if they adapt their function to the current goal. In this study, rats engaged in a naturalistic foraging task and produced defensive responses to a robotic predator. This approach allowed continuous monitoring of the neural activity as the rat switched between navigation and strategic reward procurement. By employing population-level and single-level analyses, we explored the mPFC’s role across multiple behaviors and observed a functional mode switch characterized by the co-representation of different functions in distinct spatiotemporal regions.

Results

Mix of avoidance and escape behaviors during naturalistic foraging under threat

Earlier version: The foraging arena, modified from (2018), was designed to encourage naturalistic foraging involving both approach to food and avoidance of threat. The arena was divided into two distinctive zones by a walled passage (Figure 1A): the foraging (F) and nest (N) zones. At one end of the F-zone, a robot compartment was located, separated by a motorized gate. The part of the F-zone where a rat could push its body into the robot compartment was separately named the encounter zone (E). The entrance to the E-zone was monitored by the beam sensor mounted near the gate. Inside the robot compartment, a sucrose port was situated, activated by licking, and guarded by a robotic claw (called Lobsterbot owing to its striking attack; Figure 1B).

Revised version: The foraging arena, modified from (Kimm & Choi, 2018), was designed to encourage naturalistic foraging involving both approach to food and avoidance of threat. The arena was divided into two distinctive zones by a wall with a pass-through opening in the middle (Figure 1A): the foraging (F) and nest (N) zones. At one end of the F-zone, a motorized gate controlled the access to the robot compartment, which houses a sucrose port and the predator-like Lobsterbot, named after a pair of its large, claw-like arms. The rat’s entrance to the sucrose port was detected by the beam sensor mounted near the gate (Figure 1B). A semicircular area surrounding the robot compartment was designated as encounter (E) zone where the rat’s body was situated during sucrose licking and Lobsterbot attack.

Earlier version: Behavioral procedure and apparatus. (A)

Revised version: Behavioral procedure and apparatus.

Earlier version: Schematic illustration of the arena showing major sections (N: nest zone, F: foraging zone, E: encounter zone). At the end of the E-zone, Lobsterbot (red) is situated, guarding the sucrose delivery port (green). (B) A rendered 3-D image of Lobsterbot. The sucrose port is located between the “claws”. The two red lines indicate infrared detectors, one for lick detection (short line) and the other for the entry to the E-zone (long line). (C) The experimental schedule. (D) Sample snapshots of Avoidance Withdrawal (AW) and Escape Withdrawal (EW). In a AW trial, the rat typically retracts its head ahead of time and watches the Lobsterbot attack. In a EW trial, the rat reflexively flees from the attack. (E) Example behavior data containing two consecutive trials (Trial 15 and 16). Each trial started with a reentry to the N zone which triggers gate opening. The rat leaving the N zone typically moves toward the E-zone across the F-zone. The entry to the E zone is detected by an IR beam sensor (blue shade). Within the E zone, the rat starts licking (green lines) until being attacked by Lobsterbot (red line) 3 or 6 s after the first lick. The rat shows voluntary withdrawal behavior (AW; Trial 15) or forced escape behavior (EW; Trial 16). (F) A summary of the AW trial rates for each animal during the L sessions.

Revised version: (A) Schematic illustration of the arena showing major sections (N: nest zone, F: foraging zone, E: encounter zone). At the end of the E-zone, Lobsterbot (red) is situated, guarding the sucrose delivery port (green). (B) A rendered 3-D image of Lobsterbot. The sucrose port is located between the “claws”. The two red lines indicate infrared detectors, one for lick detection (short line) and the other for the entry to the E-zone (long line). (C) The experimental schedule. (D) Sample snapshots of Avoidance Withdrawal (AW) and Escape Withdrawal (EW). In an AW trial, the rat typically retracts its head ahead of time and watches the Lobsterbot attack. In an EW trial, the rat reflexively flees from the attack. (E) Example behavior data containing two consecutive trials (Trial 15 and 16). Each trial started with a reentry to the N zone which triggers gate opening. The rat leaving the N zone typically moves toward the E-zone across the F-zone. The entry to the E zone is detected by an IR beam sensor (blue shade). Within the E zone, the rat starts licking (green lines) until being attacked by Lobsterbot (red line) 3 or 6 s after the first lick. The rat shows voluntary withdrawal behavior (AW; Trial 15) or forced escape behavior (EW; Trial 16). (F) A summary of the AW trial rates for each animal during the Losterbot sessions. Points for Lob2 of Rat2 and Rat3 are omitted because they did not approach to robot during the entire Lob2 session.

Earlier version: The entire experiment was composed of three distinct phases (Figure 1C): habituation (H), shuttling (S), and foraging under Lobsterbot threat (L). The L phase was further divided into early (L1) and late (L2) phases separated by the time of surgery. During the H phase, a food-deprived rat was habituated to the sucrose port that released a drop of sucrose solution activated by licking. Once the rat maintained a stable licking frequency, typically after 3 days, the next phase (S) started. During the S phase, the rat was trained to shuttle back and forth to procure sucrose. The robot compartment’s gate was closed 6 s from the first lick and would only reopen once the rat returned to the N-zone. The S phase lasted 3 days, followed by the introduction of the L phase. During the L phase, the robot executed a fast-striking attack, 3 s (30%) or 6 s (70%) from the first lick, varying randomly each trial to mimic natural uncertainty and prevent simple time-counting strategies. As in the S phase, the gate was automatically closed after the attack, and the rat must return to the N-zone to reopen the gate. In each trial of the L phase, the foraging rat was given a chance to respond with either an avoidance withdrawal (AW: retraction of the head before the attack) or an escape withdrawal (EW: reflexive retraction after the attack; Figure 1D). Figure 1E shows a representative sequence of responses and events in two consecutive trials. The AW was clearly distinguished by the wide gap between head withdrawal (indicated by the end of the blue shade) and the attack (red line). Conversely, the characteristic feature of the EW was the overlap between the attack and the animal’s presence in the E-zone (Video 1).

Revised version: The experiment was composed of three distinct phases (Figure 1C): habituation phase, shuttling phase, and Lobsterbot phase. The Lobsterbot phase was further divided into early and late phases separated by the time of surgery. During the habituation phase, a food-deprived rat was acclimated to the sucrose port that released sucrose solution. The lick port was activated by an IR-beam sensor, triggering the solenoid valve when the beam was interrupted. The rat gradually learned to obtain rewards by continuously licking the port. Once the rat maintained a stable licking frequency, typically after 3 days, the next phase started. During the shuttling phase, the rat was trained to shuttle back and forth to procure sucrose. The robot compartment’s gate closed 6 seconds after the first lick and reopened only when the rat returned to the N-zone. The shuttling phase lasted 3 days, followed by the introduction of the Lobsterbot phase. During this phase, the robot executed a fast-striking attack, 3 s (30%) or 6 s (70%) from the first lick, with the timing randomized across trials to mimic natural uncertainty and prevent simple time-counting strategies. As in the shuttling phase, the gate was automatically closed after the attack, and the rat must return to the N-zone to reopen the gate. In each trial of the Lobsterbot phase, the foraging rat could respond with either an avoidance withdrawal (AW: retraction of the head before the attack) or an escape withdrawal (EW: reflexive retraction after the attack; Figure 1D). Figure 1E shows a representative sequence of responses and events in two consecutive trials. The AW was clearly distinguished by a wide gap between head withdrawal (indicated by the end of the blue shade) and the attack (red line). Conversely, the characteristic feature of the EW was the overlap between the attack and the animal’s presence in the E-zone (Video 1 & 2).

Added in the revised version: We monitored all head withdrawal timepoints to assess whether rats developed a temporal strategy to differentiate between the 3-second and 6-second attacks. We found no evidence of such a strategy, as the timings of premature head withdrawals during the 6-second attack trials were evenly distributed (see Supplementary Figure S1).

Earlier version: All five rats exhibited distinctive phasic behavioral changes. Notably, there was a significant drop in approach attempts (i.e., number of trials) during the first session of the L1 phase (Lob1), followed by a stable rate of approach attempts in subsequent sessions (Supplementary Figure S1A). Other behaviors also followed similar phasic changes, including the number of licks (Supplementary Figure S1B), licks per trial (Supplementary Figure S1C), and lick latency after gate opening (Supplementary Figure S1D). To assess these behavioral changes, we compared the last session of the S phase (Sht3) and all sessions of the L1 phase (Lob1, Lob2, Lob3). The analysis showed that the number of approach attempts (repeated measure one-way analysis of variance (ANOVA); F(1.29, 5.16) = 7.18, p =.038), the number of lickings (F(1.64, 6.56) = 12.87, p =.006), and the licks per trial (F(2.21, 8.84) = 23.13, p <.001) was significantly changed during the L1 phase, while the lick latency (mixed-effects analysis; F(1.11, 3.71) = 4.06, p = .119) only showed numerical change. More specifically, approach attempts and licking were significantly decreased at the first session of the L1 phase (Sht3 vs. Lob1; t(4) = 4.53, p = .031; t(4) = 5.81, p =.013) and subsequently increased at the end of the L1 phase (Lob1 vs. Lob3; t(4) = 8.13, p = .004; t(4) = 4.84, p =.025). In summary, after the three L sessions, the rats’ sucrose consumption reverted to S phase levels (Sht3 vs. Lob3; t(4) = 1.16, p = .670) by increasing their number of approaches, compensating for a numerically decreased number of licks per trial (Sht3 vs. Lob3; number of approaches: t(4) = 6.13, p = .011; number of licks per trial: t(4) = 3.46, p = .075).

Revised version: All five rats exhibited distinctive phasic behavioral changes. Notably, there was a drop in approach attempts (i.e., number of trials) during the first session of the Lobsterbot phase (Lob1), followed by a stable rate of approach attempts in subsequent sessions after the surgery (starting from Lob4) (Supplementary Figure S2A). Other response measures, such as the number of licks (Supplementary Figure S2B), licks per trial (Supplementary Figure S2C) and lick latency after gate opening (Supplementary Figure S2D), also showed a similar phasic change. To assess these behavioral changes, we compared the last session of the shuttling phase (Sht3) and all sessions of the Lobsterbot phase before the surgery (Lob1, Lob2, Lob3). The analysis showed that the number of approach attempts (Friedman test; χ²(3)=9.240, p=.017), the number of licks (χ²(3)=12.60, p <.001), the lick latency (χ²(3)=7.800, p =.044), and the licks per trial (χ²(3)=13.08, p <.001) were significantly changed during these sessions. More specifically, number of licks and licks per trial were decreased at the first session of the Lobsterbot phase (Dunn’s multiple comparisons test; p=.021; p=.043) and subsequently increased back at the end of the early Lobsterbot phase (Lob3; p>.999; p>.999). In summary, after the three Lobsterbot sessions, the rats’ number of approach attempts, number of licks, licks per trial, and lick latency reverted to shuttling phase levels.

Earlier version: In the early sessions of the L1 phase, the rats showed passive defensive behaviors such as freezing (Blanchard & Blanchard, 1969; McClelland & Colman, 1967) and stretched attend posture (Grewal et al., 1997; Mackintosh & Grant, 1963). However, these behaviors gradually diminished toward the end of the L1 phase. Compare to the S phase, where the rats remained in the E-zone until the gate closed in 88.96% trials, they began to display a mix of AW and EW by the end of the L1 phase (AW ratio; 65.11%). Due to the transient nature of the decreased approach attempts and the various defensive behaviors after the first encounter, surgery was performed after three L sessions when the rats displayed a range of 30% to 60% AW (Figure 1F).

Revised version: In the early sessions of the Lobsterbot phase, the rats showed passive defensive behaviors such as freezing (Blanchard & Blanchard, 1969; McClelland & Colman, 1967) and stretched attend posture (Grewal et al., 1997; Mackintosh & Grant, 1963). However, these behaviors gradually diminished toward the middle of the Lobsterbot phase. Compared to the shuttling phase, during which the rats continue to lick the sucrose port until the gate is closed on almost all trials (88.96%), a mix of AW and EW (65.11% and 34.89% respectively) trials were observed during the Lobsterbot phase. Due to the transient nature of the decreased approach attempts and the various defensive behaviors after the first encounter, surgery was performed after three Lobsterbot sessions (Figure 1F).

Earlier version: Population activity in the PL and IL predicts the distance from the goal while navigating the arena

Revised version: Population activity in the mPFC predicts the distance from the Lobsterbot while navigating in the arena

Earlier version: Recording data were collected from a total of five rats implanted with a custom-made moveable tetrode drive. The tetrodes were initially implanted in the PL and advanced ventrally toward the IL at the end of each session (0.1∼0.2 mm/session). Recording locations were reversely calculated from the marking lesion that was performed after the last session. Tetrode tracks were identified by visual inspection under a microscope as indicated by black vertical lines on the matching atlas (Paxinos & Watson, 2007) as shown in Figure 2A. All tetrodes except for one spanned both the PL and the IL. After cluster cutting and preprocessing (see Materials and Methods), 632 units (57 ∼ 198 units per rat) from the PL (n = 379) and the IL (n = 253) were selected for ensemble analysis.

Revised version: Recording data were collected from a total of five rats implanted with a custom-made moveable tetrode drive. The tetrodes were initially implanted in the PL and advanced ventrally toward the IL at the end of each session (0.1∼0.2 mm/session). Recording locations were reversely calculated from the marking lesion made after the last recording session. Tetrode tracks were identified by visual inspection under a microscope as indicated by black vertical lines on the matching atlas (Paxinos & Watson, 2007) as shown in Figure 2A. All tetrodes except for one spanned both the PL and the IL. After cluster cutting and preprocessing (see Materials and Methods), 632 units (57 ∼ 198 units per rat) from the PL (n = 379) and the IL (n = 253) were selected for ensemble analysis.

Earlier version: Ensemble activity from PL and IL predicts distance from the goal. (A)

Revised version: Ensemble activity from the mPFC predicts distance from the goal.

Earlier version: Schematic diagram of electrode implantation and estimated recording site. Top: A movable 4-tetrode microdrive was initially implanted in the PL region and lowered ventrally toward the IL after every recording session. Bottom: Representative recording tracks from all five animal are superimposed over an image of a stained coronal section of the frontal brain. Histological examination of all brain sections confirmed that the electrode tracks spanned the dorsoventral axis between the PL and IL. (B) Schematics of the ensemble decoding analysis. The 4-layer deep artificial neural network (ANN) receives populational neural data during 50ms-timewindow and is trained to predict the rat’s current distance from Lobsterbot. The example data depicted in the figure is a sample recording from 20 units when the rat is at a particular distance away from Lobsterbot, indicated by the white bold line. (C) Accuracy of the location prediction. Mean Absolute Error (MAE) was calculated for the two types of decoders: one with the original dataset (Original) and another with the shuffled dataset (Shuffled). The average MAE was 16.61 cm for the Original, which was significantly smaller than that for the Shuffled and smaller than the rat’s body size. This suggests that the PL and IL might encode the spatial distance from the goal. (D) Prediction accuracy in the F-zone during outbound/inbound paths. Decoding accuracy in the F-zone was calculated separately for the outbound (from the N-zone to the E-zone) and inbound (from the E-zone to the N-zone) paths.

Revised version: (A) Schematic diagram of electrode implantation and estimated recording site. Top: A movable 4-tetrode microdrive was initially implanted in the PL region and lowered ventrally toward the IL after every recording session. Bottom: Representative recording tracks from all five animal are superimposed over an image of a stained coronal section of the frontal brain. Histological examination of all brain sections confirmed that the electrode tracks spanned the dorsoventral axis between the PL and IL. (B) Modulation of unit firing showing place-cell like activities. Units 66 and 125 exhibit fragmented place fields in all over the arena, while Units 56 and 26 display relatively large place fields surrounding particular spots such as the gates. Heat maps are calculated from z-scored spatial tuning curves. (C) Schematics of the ensemble decoding analysis. The 4-layer deep artificial neural network (ANN) receives populational neural data during 50ms-timewindow and is trained to predict the rat’s current distance from the center of the E-zone. The example data depicted in the figure is a sample recording from 20 units when the rat is at a particular distance away from the center of the E-zone, indicated by the white bold line. (D) Accuracy of the regressor. Mean Absolute Error (MAE) was calculated for the two types of regressor: one with the original dataset (Original) and another with the shuffled dataset (Shuffled). The average MAE was 16.61 cm for the Original, which was significantly smaller than that for the Shuffled and smaller than the rat’s body size. This suggests that the mPFC might encode the spatial correlates reflecting distance from the goal. (E) Prediction accuracy in the F-zone during outbound/inbound paths. Decoding accuracy in the F-zone was calculated separately for the outbound (from the N-zone to the E-zone) and inbound (from the E-zone to the N-zone) paths. The decoding accuracy remained unchanged despite the differences in motivation and perceived visual cues due to the movement direction. (F) Comparison of the regressor’s accuracy from the control experiment. When the Lobsterbot was removed from the robot compartment, reverting the task back to simple shuttling, the mPFC distance regressor’s performance significantly decreased compared to the Lobsterbot phase.

Added in the revised version: Previous studies have reported spatial correlates were observed in multiple subregions of the mPFC (Hok et al., 2005; Jung et al., 1998; Ma et al., 2023; Mashhoori et al., 2018; Sauer et al., 2022; Spellman et al., 2015). To determine whether location-specific neural activity exists at the single-cell level in our mPFC data, a traditional place map was constructed for individual neurons. Although most neurons did not show cohesive place fields across different regions in the arena, a few neurons modulated their firing rates based on the rat’s current location. However, even these neurons were not comparable to place cells in the hippocampus (O’Keefe & Dostrovsky, 1971) or grid cells in the entorhinal cortex (Hafting et al., 2005) as the place fields were patchy and irregular in some cases (Figure 2B; Units 66 and 125) or too large, spanning the entire zone rather than a discrete location within it (Units 26 and 56). The latter type of neuron has been identified in other studies (e.g., Kaefer et al., 2020).

Earlier version: Previous studies have reported spatial correlates were observed in multiple subregions of the mPFC (Hok et al., 2005; Jung et al., 1998; Ma et al., 2023; Mashhoori et al., 2018; Sauer et al., 2022; Spellman et al., 2015). To determine whether the mPFC units had spatial correlates, we performed a regression analysis for the entire arena (N-, F-, and E-zones). The inputs to the regressor were the simultaneously recorded neural activity and the predicted output was the distance from the goal location computed from the rat’s head position. Among machine-learning algorithms, we implemented a four-layer artificial neural network (ANN; see Materials and Methods for a detailed structure), as the ANN achieves significantly fewer decoding errors compared to other algorithms (Mashhoori et al., 2018). More specifically, to decode spatial information from neural data, recorded neural activity was segmented into 50-ms bins and used as input; the distance from the center of the robot was used as the output (Figure 2B).

Revised version: Since we did not find neural activity that can be decoded into accurate spatial information based on individual activity, we performed machine-learning (ML) analysis targeting all recorded units during the given session. Among different ML algorithms, we implemented a four-layer artificial neural network regressor (ANN; see Materials and Methods for a detailed structure), as the ANN achieves significantly lower decoding errors (Supplementary Figure S3) and has robustness to hyperparameter changes (Glaser et al., 2020). The inputs to the regressor were the simultaneously recorded neural activity which was segmented into 50-ms bins and the predicted output was the distance between the rat’s head and the center of the E-zone where reward and threat co-exist (Figure 2C).

Earlier version: The analysis of trained ANNs revealed that the distance from the rat to the robot can be decoded from the mPFC ensemble activity. Compared to the control regressor trained on a shuffled dataset, the regressor with original data showed significantly lower errors (paired t-test, t(39) = 10.47, p <.001), indicating that the population activity in the PL and the IL contain spatial information (Figure 2C, Video 2). The overall accuracy defined in the form of the mean absolute error (MAE) was 16.61 (+-3.67) cm, which was slightly higher than values from previous reports employing different spatial tasks (Ma et al., 2023; Mashhoori et al., 2018). In the current study, an average of 15.8 neurons, a number lower than that of previous studies, was included in a given ensemble in our analysis. Owing to the nature of the machine-learning algorithm, we expected that accuracy would be higher with a greater number of simultaneously recorded neurons.

Revised version: The analysis of trained ANNs revealed that the distance from the rat to the center of the E-zone can be decoded from the mPFC ensemble activity. Compared to the control regressor trained on a shuffled dataset, the regressor with original data showed significantly lower errors (paired t-test, t(39) = 10.47, p <.001). However, we did not observe a difference in decoding accuracy between the PL and IL ensembles, and there were no significant interactions between regressor type (shuffled vs. original) and regions (mixed-effects model; regions: p=.996; interaction: p=.782). These results indicate that the population activity in both the PL and IL contains spatial information (Figure 2D, Video 3). The overall accuracy defined in the form of the mean absolute error (MAE) was 16.61 (± 3.67) cm, which was slightly higher than values from previous reports employing different spatial tasks (Ma et al., 2023; Mashhoori et al., 2018). In the current study, an average of 15.8 neurons, a number lower than that of previous studies, was included in a given ensemble in our analysis. Considering the nature of the machine-learning algorithm, we expected that accuracy would be higher with a greater number of simultaneously recorded neurons.

Earlier version: To determine whether the accuracy of the spatial encoding varied depending on the direction of movement, we compared the decoding accuracy of the regressor for outbound (from the N- to E-zone) vs. inbound (from the E- to N-zone) navigation within the F-zone. There was no significant difference in decoding accuracy between outbound vs. inbound trips (paired t-test; t(39) = 1.52, p =.136), indicating that the stability of spatial encoding was maintained regardless of the moving direction or perceived context (Figure 2D). In summary, these data suggest that precise spatial encoding could be extracted using the neural ensemble in the mPFC.

Revised version: To determine whether the accuracy of the regressor varied depending on the direction of movement, we compared the decoding accuracy of the regressor for outbound (from the N-to E-zone) vs. inbound (from the E- to N-zone) navigation within the F-zone. There was no significant difference in decoding accuracy between outbound vs. inbound trips (paired t-test; t(39) = 1.52, p =.136), indicating that the stability of spatial encoding was maintained regardless of the moving direction or perceived context (Figure 2E).

Added in the revised version: Previous analyses indicated that the distance regressor performed robustly regardless of movement direction, but there is a possibility that the decoder detects visual cues or behaviors specific to the E-zone. For example, neural activity related to Lobsterbot confrontation or licking behavior might be used by the regressor to decode distance. To rule out this possibility, we analyzed a subset of data collected when the compartment door was closed, preventing visual access to the Lobsterbot and sucrose port and limiting active foraging behavior. The regressor trained on this subset still decoded distance with a MAE of 12.14 (± 3.046) cm (paired t-test; t(39) = 12.17, p <.001). Notably, the regressor’s performance was significantly higher with this subset than with the full dataset (paired t-test; t(39) = 9.895, p <.001). In summary, these data suggest that precise spatial information could be extracted using the neural ensemble in the mPFC.

Removed in the revised version: The decoding accuracy remained unchanged despite the differences in motivation and perceived visual cues due to the movement direction.

Removing the Lobsterbot decreases spatial decoding accuracy

Earlier version: Earlier studies have suggested that PL neurons do not exhibit any spatially correlated firing patterns in an empty round apparatus (Poucet, 1997). This absence of spatial correlates may result from a lack of goal-oriented navigation behavior. To evaluate how the Lobsterbot’s presence impacted spatial decoding accuracy, additional single-unit recording experiments (Ctrl-Exp) were conducted and compared with the previous results (Lob-Exp). In this experiment, three rats followed the same protocol until the microdrive implant surgery. After the surgery, unlike the Lob-Exp group, the Ctrl-Exp group returned to the S phase, where the Lobsterbot was removed. During a total of nine recording sessions, 112 units were identified and ANN regressors were built and trained using these data. Statistical tests indicated that the original dataset had significantly lower decoding errors (20.48 ±2.31 cm) compared to the shuffled dataset (22.61 ±1.95 cm; paired t-test; t(8) = 6.20, p <.001). Nevertheless, when the error was compared with that of Lob-Exp, the Ctrl-Exp group showed a higher decoding errors (unpaired t-test; t(47) = 3.02, p =.004). This did not result from differences in activity levels, as there was no significant difference in travel distance between the Lob-Exp and Ctrl-Exp groups (unpaired t-test; t(47) = 0.07, p =.941).

Revised version: Earlier studies have suggested that PL neurons do not exhibit any spatially correlated firing patterns in an empty round apparatus (Poucet, 1997). This absence of spatial correlates may result from a lack of complex goal-oriented navigation behavior, which requires deliberate planning to acquire more rewards and avoid potential threats. To evaluate how the Lobsterbot’s presence impacted spatial decoding accuracy, additional single-unit recording experiments (Ctrl-Exp) were conducted and compared with the previous results (Lob-Exp). In this experiment, three rats followed the same protocol until the microdrive implant surgery. After the surgery, unlike the Lob-Exp group, the Ctrl-Exp group returned to the shuttling phase, during which the Lobsterbot was removed. With this protocol, both groups experienced sessions with the Lobsterbot, but the Ctrl-Exp group’s task became less complex, as it was reduced to mere reward collection. During a total of nine recording sessions, 112 units were identified and ANN regressors were built and trained using these data. Statistical tests indicated that the original dataset had significantly lower decoding errors (20.48 ±2.31 cm) compared to the shuffled dataset (22.61 ±1.95 cm; paired t-test; t(8) = 6.20, p <.001). Nevertheless, when the error was compared with that of Lob-Exp, the Ctrl-Exp group showed a higher decoding errors (Figure 2F; unpaired t-test; t(47) = 3.02, p =.004). This did not result from differences in activity levels, as there was no significant difference in total travel distance between the Lob-Exp and Ctrl-Exp groups (unpaired t-test; t(47) = 0.07, p =.941).

Another factor potentially contributing to the increase in decoding errors could be differences in recording quality across experiments. To discriminate the effects, a generalized linear model (GLM) was constructed with predictors of 1) the Lobsterbot’s presence, 2) the number of units recorded, and 3) the recording site (PL vs. IL) on decoding error. The results showed that the GLM predicted the outcome variable with significant accuracy (F(3,45) = 14.06, p <.001), explaining 48.38% of the variance in decoding error. The regression coefficients were as follows: presence of the Lobsterbot (2.76, standard error [SE] = 1.11, t = 2.49, p =.016), number of recorded cells (−0.43, SE =.08, t = 5.29, p <.001), and recording location (0.91, SE =.92, p =.329). The results indicated that the number of recorded cells significantly influenced decoding accuracy. Most importantly, the presence of the Lobsterbot significantly decreased the overall decoding error, even after accounting for the number of recorded cells. In summary, the removal of the Lobsterbot decreased spatial decoding accuracy.

Added in the revised version: Non-navigational behavior reduces the accuracy of decoded location

Removed in the revised version: Excluding non-navigational behavior (stereotypic behavior) enhances spatial accuracy of the population decoding in the N-zone

Removed in the revised version: Although the initial analysis showed that the population activity encoded the distance of the rat to the goal reasonably well, its spatial precision was unevenly distributed throughout the arena (Figure 3A). The magnitude of error was highest in the N-zone, followed by the E-zone, and lowest in the F-zone. A one-way repeated measures ANOVA revealed a significant main effect of the zone (F(1.68, 65.52) = 40.36, p <.001). Post hoc analysis showed that the MAE in the F-zone (13.86 ± 2.56 cm) was significantly lower than that in the N-zone (22.72 ± 6.38 cm; t(39) = 10.69; p <.001) and the E-zone (19.98 ± 7.30 cm; t(39) = 6.34; p <.001).

Removed in the revised version: Spatial encoding is disrupted by non-navigational behaviors (NNBs). (A)

Removed in the revised version: Spatial distribution of the prediction accuracy. The heatmap indicates MAE of a fully trained ANN, superimposed over the entire foraging arena. The prediction accuracies were lower in the N- and E-zone than that in the F-zone. (B) Mean prediction accuracy by the zones. The MAE in F-zone was significantly lower than the other zones. Error bars represent the SEM. (C) Examples of the NNBs in the N-zone. The top three snapshots depict grooming, rearing and sniffing. The bottom three snapshots show typical goal-directed navigational movements. (D) Comparison of location-decoding errors (N-zone) during navigational vs. NNBs. The error was significantly larger when the rat was engaged in NNBs within the N-zone. (E) Comparison of location-decoding errors of regressors trained with and without NNBs. The overall decoding error was significantly smaller when the ANN was trained without the data during NNBs.

Earlier version: However, there was no significant difference between the N- and E-zones (t(39) = 2.29; p =.081; Figure 3B). To correct for the unequal distribution of location visits (more visits to the F-than to other zones), the regressor was trained using a subset of the original data, which was equalized for the data size per distance range (see Materials and Methods). Despite the correction, there was a significant main effect of the zone (F(1.16, 45.43) = 119.2, p <.001) and the post hoc results showed that the MAEs in the N-zone (19.52 ± 4.46 cm; t(39) = 10.45; p <.001) and the E-zone (26.13 ± 7.57 cm; t(39) = 11.40; p <.001) had a significantly higher errors when compared to the F-zone (14.10 ± 1.64 cm).

Revised version: Although initial analysis showed that the distance of the rat to the center of the E-zone could be decoded with reasonable accuracy, the spatial precision was unevenly distributed throughout the arena (Figure 3A). The magnitude of error was highest in the N-zone, followed by the E-zone, and lowest in the F-zone. A one-way repeated measures ANOVA revealed a significant main effect of the zone (F(1.68, 65.52) = 40.36, p <.001). Post hoc analysis showed that the MAE in the F-zone (13.86 ± 2.56 cm) was significantly lower than that in the N-zone (22.72 ± 6.38 cm; t(39) = 10.69; p <.001) and the E-zone (19.98 ± 7.30 cm; t(39) = 6.34; p <.001). However, there was no significant difference between the N- and E-zones (t(39) = 2.29; p =.081; Figure 3B). To correct for the unequal distribution of location visits (more visits to the F-than to other zones), the regressor was trained using a subset of the original data, which was equalized for the data size per distance range (see Materials and Methods). Despite the correction, there was a significant main effect of the zone (F(1.16, 45.43) = 119.2, p <.001) and the post hoc results showed that the MAEs in the N-zone (19.52 ± 4.46 cm; t(39) = 10.45; p <.001) and the E-zone (26.13 ± 7.57 cm; t(39) = 11.40; p <.001) had a significantly higher errors when compared to the F-zone (14.10 ± 1.64 cm).

Added in the revised version: Spatial encoding is disrupted by non-navigational behaviors

Added in the revised version: (A) Spatial distribution of the prediction accuracy. The heatmap indicates MAE of a fully trained ANN, superimposed over the entire foraging arena. The prediction accuracies were lower in the N- and E-zone than that in the F-zone. (B) Mean prediction accuracy by the zones. The MAE in F-zone was significantly lower than the other zones. Error bars represent the SEM. (C) Examples of the non-navigational behaviors in the N-zone. The top three snapshots depict grooming, rearing and sniffing. The bottom three snapshots show typical goal-directed navigational movements. (D) Comparison of decoding errors (N-zone) during navigational vs. non-navigational behaviors. The error was significantly larger when the rat was engaged in non-navigational behaviors within the N-zone. (E) Comparison between regressors trained with and without non-navigational behaviors. The overall decoding error was significantly smaller when the regressor was trained without the data during non-navigational behaviors.

Earlier version: In the N-zone, the rats engaged in stereotypic behaviors such as grooming, climbing, and rearing, during which they stopped locomotive behavior (Figure 3C). To determine whether these non-navigational behaviors resulted in higher decoding errors, these behaviors inside the N-zone is manually labelled and the MAEs during navigational vs. non-navigational behaviors were compared. Due to the imbalanced nature of the dataset resulting from varying preference for non-navigational behaviors, data from 21 sessions were included, each containing at least 20% of one type of behavior. The paired t-test showed a significant difference between the MAEs during the non-navigational and navigational behaviors in the N-zone (paired t-test; t(20) = 9.381, p <.001; Figure 3D). Moreover, excluding the non-navigational data of the N-zone from the original dataset improved the overall accuracy of the regressor (paired t-test; t(20) = 8.562, p <.001; Figure 3E). These data indicate that erroneous spatial encoding in the N-zone might have resulted from the disruption induced by non-navigational behaviors.

Revised version: In the N-zone, the rats engaged in stereotypic behaviors such as grooming, climbing, and rearing, during which they stopped locomotive behavior (Figure 3C). To determine whether these non-navigational behaviors resulted in higher decoding errors, these behaviors inside the N-zone is manually labelled and the MAEs during navigational vs. non-navigational behaviors were compared. Due to the imbalanced nature of the dataset resulting from varying preference for non-navigational behaviors, data from 21 sessions were included, each containing at least 20% of one type of behavior. The paired t-test showed a significant difference between the MAEs during the non-navigational and navigational behaviors in the N-zone (paired t-test; t(20) = 9.381, p <.001; Figure 3D). Moreover, excluding the non-navigational data of the N-zone from the original dataset improved the overall accuracy of the regressor (paired t-test; t(20) = 8.562, p <.001; Figure 3E). These data indicate that erroneous decoding result in the N-zone might have resulted from the disruption induced by non-navigational behaviors.

Earlier version: Multi-dimensional analysis reveals a different functional mode in the E-zone

Revised version: Population analysis reveals a different functional mode in the E-zone

Earlier version: To assess the functional coherence of population activity (Durstewitz et al., 2010) within and across different zones, principal component analysis (PCA) was performed on the ensemble data collected throughout the session. Figure 4A illustrates the first two principal components of the transformed neural data from a representative recording session (number of units = 19). First, within-zone and between-zone distances in population neural activity were compared to determine whether activity within the same zone exhibited similar properties, distinct from those of other zones A paired t-test revealed a significant difference between distances (t(119) = 10.46, p <.001), confirming that the neural activity within a given zone was distinct from those of other zones.

Revised version: To assess the functional coherence of population activity (Durstewitz et al., 2010) within and across different zones, principal component analysis (PCA) was performed on the ensemble data collected throughout the session. Figure 4A illustrates the first two principal components of the transformed neural data from a representative recording session (number of units = 19). First, within- and between-zone distances in population neural activity were compared to determine whether activity within the same zone exhibited similar properties, distinct from those of other zones A paired t-test revealed a significant difference between distances (t(119) = 10.46, p <.001), confirming that the neural activity within a given zone was distinct from those of other zones.

Earlier version: PCA results reveals distinctive population activity in the E-zone (A)

Revised version: PCA results reveals distinctive population activity in the E-zone

Earlier version: Representative recording session depicting the first two dimensions of the PCA result. Populational neural activities are projected onto a virtual space. Each dot represents 50 ms-long neural activity from multiple units. The color of the dots indicates the rat’s location during the corresponding neural activity. Diamonds represent the centroids of neural representations for each zone. To visually emphasize each cluster, data points close to centroids are selectively plotted. (B) Distances between each centroid pairs from all recording sessions. The centroid of the E-zone is distinctly positioned compared to the centroids of the other two zones, indicating a unique neural state within the E-zone. The triangle above the bar plot represents the relative distance between the centroids of each zone’s neural ensemble activity. Longer edges signify greater dissimilarity between the neural ensemble activities of two zones. Error bars represent the SEM.

Revised version: (A) Representative recording session depicting the first two dimensions of the PCA result. Populational neural activities are projected onto a virtual space. Each dot represents 50 ms-long neural activity from multiple units. The color of the dots indicates the rat’s location during the corresponding neural activity. Diamonds represent the centroids of neural representations for each zone. To visually emphasize each cluster, data points close to centroids are selectively plotted. (B) Distances between each centroid pairs from all recording sessions. The centroid of the E-zone is distinctly positioned compared to the centroids of the other two zones, indicating a unique neural state within the E-zone. The triangle above the bar plot represents the relative distance between the centroids of each zone’s neural ensemble activity. Longer edges signify greater dissimilarity between the neural ensemble activities of two zones. Error bars represent the SEM.

Earlier version: To quantitatively measure functional resemblance among ensemble activity within different zones, we measured the distance between pairs of centroids projected onto the latent space (N-F, N-E, and F-E). A one-way repeated measure ANOVA showed a significant difference between centroid pairs (F(1.61, 62.64) = 51.99, p <.001). Further post hoc analysis revealed a significant difference between N-F and N-E pairs (t(39) = 9.96, p <.001), N-F and F-E pairs (t(39) = 2.55, p =.044), and N-E and F-E pairs (t(39) = 9.13, p <.001), suggesting an ordinal similarity with N-F (D =.72) being most similar, followed by F-E (D = .94) and N-E (D = 1.44) (Figure 4B).

Revised version: To quantitatively measure functional resemblance among ensemble activity within different zones, we measured the distance between pairs of centroids projected onto the latent space (N-F, N-E, and F-E). A Friedman test showed a significant difference between centroid pairs (χ²(2)=51.05, p <.001). Further post hoc analysis revealed a significant difference between N-F and N-E pairs (p <.001) and N-E and F-E pairs (p <.001) but not in N-F and F-E pairs (p =.353), suggesting an ordinal similarity with N-F (D =.72) being most similar, followed by F-E (D = .94) and N-E (D = 1.44) (Figure 4 Figure 4B).

Earlier version: Not surprisingly, the firing pattern in the E-zone was the most dissimilar, reflecting a functional state different from those in other zones. In contrast, the distance between centroids of the N- and F-zones was relatively small, indicating that the overlap between the neural subspaces in the two zones was not negligible. In summary, populational activity inside the E-zone showed distinctive activity patterns unlike those of active navigation.

Revised version: Not surprisingly, the firing pattern in the E-zone was the most dissimilar, reflecting a functional state different from those in other zones. In contrast, the distance between centroids of the N- and F-zones was relatively small, indicating that the overlap between the neural subspaces in the two zones was not negligible.

Added in the revised version: Although PCA can capture population differences between neural ensembles in the zones, a small subset of neurons highly tuned to salient sensory cues could solely drive this separation. Particularly in the E-zone, which encompasses two significant events—head-entry and head-withdrawal—neurons responsive to the delivery of rewards or attacks from the robot could dominate the principal components. To minimize this influence, we defined “critical event times” as ±1 second from head-entry and head-withdrawal, excluded neural data during these periods, and then recalculated the principal components. The results still demonstrated the same sequential ensemble differences (F(1.57, 61.34) = 73.56, p <.001), indicating that the population differences are not driven by events specific to any zone (Supplementary Figure S4). In summary, populational activity inside the E-zone showed distinctive activity patterns unlike those of active navigation.

Earlier version: Hierarchical clustering confirms functionally distinctive sub-populations in the E-zone

Revised version: Hierarchical clustering identifies functionally distinctive sub-populations in the E-zone

Earlier version: In the E-zone, navigational behavior was mostly replaced by approach attempts and avoidance. To characterize the temporal dynamics of single-unit activity in the E-zone, we examined the modulation of firing rate to the two most critical behaviors; the head entry (HE) and the head withdrawal (HW). The temporal dynamic of neuronal firing relative to these two behavioral events was visualized using color-coded raster plots and peri-event time histograms (Figure 5A & 5C), featuring spiking activities from all event-responsive units (see Materials and Methods). Overall, the units showed modulated firing with varying degrees of temporal relations to the two events. Most strikingly, overall modulation patterns were clearly different between HE and HW.

Revised version: In the E-zone, navigational behavior was mostly replaced by approach attempts and avoidance behaviors. To characterize the temporal dynamics of single-unit activity in the E-zone, we examined the modulation of firing rate to the two most critical behaviors; the head-entry and the head-withdrawal. The temporal dynamic of neuronal firing relative to these two behavioral events was visualized using color-coded raster plots and peri-event time histograms (Figure 5A & 5C), featuring spiking activities from all behavior-responsive units (see Materials and Methods). Overall, the units showed modulated firing with varying degrees of temporal relations to the two events. Most strikingly, overall modulation patterns were clearly different between head-entry and head-withdrawal.

Added in the revised version: Multiple subpopulations in the mPFC react differently to head entry and head withdrawal.

Removed in the revised version: Peri-Event Time Histogram (PETH) of all units -responsive to the head entry (HE), aligned by the peak firing rate and the averaged activity. (A)

Earlier version: Top: The PETH for each unit is based on the Z-score of which magnitude is color-coded (temperature bar on the right). Bottom: The red vertical lines mark the timing of the HE. The peak latency of each unit varies from as early as 2 s before and to 1∼2 s after the HE. (B) Functional segregation of the recorded units. Top and middle: Two sub-populations of units based on hierarchical cell clustering analysis. Bottom: The averaged activity for each sub-population. (C) PETH of all units responsive to the head withdrawal (HW) and the averaged activity. (D) Functional segregation of the recorded units. Top and middle: Three sub-populations of units based on hierarchical cell clustering analysis. Bottom: The averaged activity for each sub-population.

Revised version: (A) Top: The PETH of head entry-responsive units is color-coded based on the Z-score of activity. Bottom: The red vertical lines mark the timing of the head-entry. The peak latency of each unit varies from as early as 2 s before and to 1∼2 s after the head-entry. (B) Functional segregation of all recorded units. Top and middle: Two sub-populations of units based on hierarchical cell clustering analysis. Bottom: The averaged activity for each sub-population. (C) The PETH of head withdrawal-responsive units is color-coded based on the Z-score of activity. (D) Functional segregation of all recorded units. Top and middle: Three sub-populations of units based on hierarchical cell clustering analysis. Bottom: The averaged activity for each sub-population.

Earlier version: For the HE, a large portion of the units increased the firing before the event (Figure 5A, top). The pattern was clearly present when all neuronal spike activities were collectively analyzed (Figure 5A, bottom). The mean activity showed a peak approximately 150 ms before the HE, with the peak exhibiting substantial width (FWHM of 200 ms), suggesting the presence of varied neuron populations differing in peak position. In contrast, this anticipatory modulation was not clearly identified when the activities were aligned to the HW. Instead, there was a strong temporal aggregation of activity around the peri-event period (Figure 5C, top) within the narrow range. This was evident in the collapsed histogram as a single sharp peak at the time of the HW (peak at 0 ms, FWHM of 75 ms; Figure 5C, bottom).

Revised version: For the head-entry, a large portion of the units increased the firing before the event (Figure 5A, top). The pattern was clearly present when all neuronal spike activities were collectively analyzed (Figure 5A, bottom). The mean activity showed a peak approximately 150 ms before the head-entry, with the peak exhibiting substantial width (FWHM of 200 ms), suggesting the presence of varied neuron populations differing in peak position. In contrast, this anticipatory modulation was not clearly identified when the activities were aligned to the head-withdrawal. Instead, there was a strong temporal aggregation of activity around the peri-event period (Figure 5C, top) within the narrow range. This was evident in the collapsed histogram as a single sharp peak at the time of the head-withdrawal (peak at 0 ms, FWHM of 75 ms; Figure 5C, bottom).

Earlier version: To categorize recorded units into multiple sub-populations based on their functional connectivity in reference to the key events (HE and HW), the unsupervised hierarchical clustering method was applied to all recorded units (Boehlen et al., 2016; Di Miceli et al., 2020). We identified two sub-populations for the HE (Figure 5B). The first sub-population (HE1: n = 366, 57.91%) showed an anticipatory increase with a peak located at 350 ms before and a negative modulation after entry and remained low for an extended period. The other sub-population (HE2: n = 189, 29.91%) showed a more concentrated positive modulation at around the time of entry. For the HW, three sub-populations were identified (Figure 5D). The first (HW1: n = 380, 60.13%) showed a complex modulation pattern: lower than the baseline in the beginning, followed by a peak firing timed at the HW. The second (HW2: n = 147, 23.26%) maintained increased activity until the onset of HW and returned to zero afterward. The last sub-population, HW3, which consisted of a relatively small number of units (n = 62, 9.81%), showed an inhibitory response at the time of HW.

Revised version: To categorize recorded units into multiple sub-populations based on their functional connectivity in reference to the key events, the unsupervised hierarchical clustering method was applied to all recorded units (Boehlen et al., 2016; Di Miceli et al., 2020). We identified two sub-populations when we classified neurons according to their activity around the head-entry (Figure 5B). The first sub-population (HE1: n = 366, 57.91%) showed an anticipatory increase with a peak located at 350 ms before and a negative modulation after entry and remained low for an extended period. The other sub-population (HE2: n = 189, 29.91%) showed a more concentrated positive modulation at around the time of entry. When we used the head withdrawal as criterion, three sub-populations were identified (Figure 5D). The first (HW1: n = 380, 60.13%) showed a complex modulation pattern: lower than the baseline in the beginning, followed by a peak firing timed at the head-withdrawal. The second (HW2: n = 147, 23.26%) maintained increased activity until the onset of head-withdrawal and returned to zero afterward. The last sub-population, HW3, which consisted of a relatively small number of units (n = 62, 9.81%), showed an inhibitory response at the time of head-withdrawal.

Earlier version: Well-trained rats typically run across the F-zone at high velocity and arrive at the sucrose port (HE) in the E-zone, stopping abruptly. Moreover, the HW is always aligned with the rats’ movements, including both subtle (AW) and reflexive (EW) movement. Consequently, it is plausible that the sudden movement might result sudden change in neural activity around the time of HE and/or HW. To exclude this possibility, a “run-and-stop” event was defined and correlated with the neural data. This event was defined as the transitional moment between a directional locomotion of 2 s or longer and a subsequent pause of 1 s or longer (see Materials and Methods). Unit activities were aligned around this run-and-stop event, and the mean activity of each sub-population was plotted for visual inspection. As shown in the Supplementary Figure S2, there was no significant neural modulation detectable around the run-and-stop events. These findings suggest that the neural modulation occurring around the HE and HW is not attributable to changes in velocity of movement.

Revised version: Well-trained rats typically run across the F-zone at high velocity and arrive at the sucrose port in the E-zone, stopping abruptly. Moreover, the head-entry is always aligned with the rats’ movements, including both subtle (AW) and reflexive (EW) movement. Consequently, it is plausible that the sudden movement might result sudden change in neural activity around the time of head-entry and/or head-withdrawal. To exclude this possibility, a “run-and-stop” event was correlated with the neural data. The event was defined as the transitional moment between a directional locomotion of 2 s or longer and a subsequent pause of 1 s or longer (see Materials and Methods). Unit activities were aligned around this run- and-stop event, and the mean activity of each sub-population was plotted for visual inspection. As shown in the Supplementary Figure S5, there was no detectable neural modulation around the run-and-stop events. The lack of significant modulation at the time of action suggests that the neural modulation occurring around the head-entry and head-withdrawal is not a simple motoric response.

Earlier version: Interestingly, a significant categorical overlap between HE and HW neurons was found. In the HW1 group, 78.68% (n = 299) of units were classified as HE1. In the HW2 group, 63.95% (n = 94) of units were from HE2 (Supplementary Figure S3A). Based on this, we further characterized the temporal dynamics of spike activity in relation to the two behavioral events using the combined categories: Type 1 (HE1-HW1) and Type 2 (HE2-HW2). These two subpopulations demonstrated distinctive firing patterns within the E-zone (Supplementary Figure S3B). Specifically, Type 1 neurons exhibited a rapid decline in activity following a short peak 200 ms before the HE, whereas Type 2 neurons displayed a sustained high peak activity (z ≍ 2) from the onset of the HE until the HW. Taken together, these findings provide support for the existence of at least two distinct subpopulations in the PL and IL that may be involved in different cognitive functions within the E-zone.

Revised version: An interesting overlap was found between sub-populations characterized by head-entry and head-withdrawal. In the HW1 group, 78.68% (n = 299) of units were classified as HE1. On the other hand, 63.95% (n = 94) of the HW2 units were from HE2 (Supplementary Figure S6A). Based on this, we further characterized the temporal dynamics of spike activity in relation to the two behavioral events using the combined categories: Type 1 (HE1-HW1) and Type 2 (HE2-HW2). These two subpopulations demonstrated distinctive firing patterns within the E-zone (Supplementary Figure S6B & S7). Specifically, Type 1 neurons exhibited a rapid decline in activity following a short peak 200 ms before the head-entry, whereas Type 2 neurons displayed a sustained high peak activity (z ≍ 2) from the onset of the head-entry until the head-withdrawal.

Added in the revised version: We also examined whether there is a significant difference between the PL and IL in the proportion of Type 1 and Type 2 neurons. In the PL, among 379 recorded units, 143 units (37.73%) were labeled as Type 1, and 75 units (19.79%) were labeled as Type 2. In contrast, in the IL, 156 units (61.66%) and 19 units (7.51%) of 253 recorded units were labeled as Type 1 and Type 2, respectively. A Chi-square analysis revealed that the PL contains a significantly higher proportion of Type 2 neurons (χ²(1, 632) = 34.85, p < .001), while the IL contains a significantly higher proportion of Type 1 neurons compared to the other region (χ²(1, 632) = 18.07, p < .001). Taken together, these findings provide support for the existence of at least two distinct subpopulations in the PL and IL that may be involved in different cognitive functions within the E-zone.

Population activities in the E-zone encode success and failure of avoidance with high accuracy

Earlier version: The previous analysis confirmed that a significant number of neurons modulated their activity before or around the time of approach toward and withdrawal from the E-zone. We then further analyzed the withdrawal-related activity to determine their functional contributions to avoidance decisions. If the recorded neurons in the mPFC encode avoidance decisions, AW and EW may be distinguishable from decoded population activity. A naïve Bayesian classifier (NBC, see Materials and Methods) was trained to distinguish AW from EW using 2-s neural ensemble data around the onset of the event (Figure 6A). We corrected the accuracy against the sample imbalance owing to varying degrees of the AW/EW response ratio across different sessions. Our analysis showed that the classifier based on the original dataset achieved a significantly higher prediction accuracy (72.71%) than a theoretical chance-level prediction (50%) or a shuffled dataset (51.62%) (paired t-test; t(39) = 10.06, p <.001) in distinguishing AW from EW (Figure 6B).

Revised version: The previous analysis confirmed that a significant number of neurons modulated their activity before or around the time of approach toward and withdrawal from the E-zone. We then further analyzed the correlation between neural activity and success and failure of avoidance behavior. If the mPFC neurons participate in the avoidance decisions, avoidance withdrawal (AW; withdrawal before the attack) and escape withdrawal (EW; withdrawal after the attack) may be distinguishable from decoded population activity even prior to motor execution. A naïve Bayesian classifier (see Materials and Methods) was trained to distinguish AW from EW using 2-s neural ensemble data around the onset of the event (Figure 6A). We corrected the accuracy against the sample imbalance owing to varying degrees of the AW/EW response ratio across different sessions. Our analysis showed that the classifier based on the original dataset achieved a significantly higher prediction accuracy (72.71%) than a theoretical chance-level prediction (50%) or a shuffled dataset (51.62%) (paired t-test; t(39) = 10.06, p <.001) in distinguishing AW from EW (Figure 6B). Furthermore, we analyzed whether there is a difference in prediction accuracy between sessions with different recorded regions, the PL and the IL. A repeated two-way ANOVA revealed no significant difference between recorded regions, nor any interaction (regions: F(1, 38) = 0.1828, p = 0.671; interaction: F(1, 38) = 0.1614, p = 0.690).

Earlier version: Neural ensemble activity predicts failure and success of avoidance response. (A)

Revised version: Neural ensemble activity predicts failure and success of avoidance response.

Earlier version: Schematics of the event decoding analysis. The Naïve Bayesian decoders (NBCs) are trained with 2 s window of neural activity to discriminate avoidance or escape on every trial (AW/EW classifier). The grayscale image depicts an example firing pattern of 17 units on a given trial, arranged to the onset of the withdrawal response. The decoder classifies whether this trial is AW or EW based on this data. (B) Accuracy of the NBC. The decoding accuracy of the classifier was significantly higher than that from the shuffled data. (C) Temporal characteristics of prediction accuracy of the NBC. Prediction accuracy was significantly higher at the time points as early as 5-7 s before the HW. (D) Class discrimination index by the two sub populations of neurons. The class discrimination index indicates that the Type 2 neurons showed a significant discriminatory power towards AW. Neurons in the Type 2 and the Others group did not exhibit significant discriminatory power.

Revised version: (A) Schematics of the event decoding analysis. The Naïve Bayesian decoders are trained with 2 s window of neural activity to discriminate avoidance or escape on every trial (AW/EW classifier). The grayscale image depicts an example firing pattern of 17 units on a given trial, arranged to the onset of the withdrawal response. The decoder classifies whether this trial is AW or EW based on this data. (B) Accuracy of the naïve Bayesian classifier. The decoding accuracy of the classifier was significantly higher than that from the shuffled data. (C) Temporal characteristics of prediction accuracy of the naïve Bayesian classifier. Prediction accuracy was significantly higher at the time points as early as 5-7 s before the head-withdrawal. (D) Class discrimination index by the two sub populations of neurons. The class discrimination index indicates that the Type 2 neurons showed a significant discriminatory power towards AW. Neurons in the Type 2 and the Others group did not exhibit significant discriminatory power.

Earlier version: Because the behavioral reactions following the two events were quite different (bigger reactions during the EW as if the animal failed to predict the attack), it is not clear whether the NBC discriminated the two types of responses based on the post-event population activity, which would have been a reflection of the different reactions in the two responses rather than an anticipatory or value-driven encoding. To clarify this issue, we trained NBCs using 2-s ensemble activity at varying time points before the execution of the HW response. Multiple t-tests between original and shuffled datasets revealed that AW/EW discrimination accuracy began to increase above the chance level 6 s before the actual response (multiple paired t-tests with Šídák-Bonerroni correction; -31 – -29: t(39) = .24, p > .999; -21 – -19: t(39) = .57, p > .999; -11 – -9: t(39) = 1.62, p = .701; -7 – -5: t(39) = 5.381, p < .001; -6 – -4: t(39) = 5.328, p < .001; -5 – -3: t(39) = 4.610, p < .001; -4 – -2: t(39) = 4.631, p < .001; -3 – -1: t(39) = 4.544, p = .008; -2 – 0: t(39) = 4.544, p < .001; -1 – +1: t(39) = 8.505, p < .001; Figure 6C).

Revised version: Because the behavioral reactions following the two events were quite different (bigger reactions during the EW as if the animal failed to predict the attack), it is not clear whether the naïve Bayesian classifier discriminated the two types of responses based on the post-event population activity, which would have been a reflection of the different reactions in the two responses rather than an anticipatory or value-driven encoding. To clarify this issue, we trained naïve Bayesian classifiers using 2-s ensemble activity at varying time points before the execution of the head-withdrawal response. Multiple t-tests between original and shuffled datasets revealed that AW/EW discrimination accuracy began to increase above the chance level 6 s before the actual response and remain significant at all time points (multiple paired t-tests with Šídák-Boferroni correction; ps < .01; Figure 6C, Supplementary Data 1).

Earlier version: Based on the previous clustering analysis that identified functionally cohesive sub-populations of neurons (Types 1 and 2), we further investigated the predictive encoding by the ensemble activity of the two subpopulations. The NBC uses class likelihood, P(unit|AW) and P(unit|EW), to classify a particular type of withdrawal response. By comparing the class likelihoods, we can assess the ability of these subgroups to distinguish between AW and EW. This ability is quantified by a ‘Class Discrimination Index,’ calculated as a log odds ratio. This index reflects the likelihood that a given feature set is associated with a particular type of withdrawal response (AW or EW). To test whether the unit activity right before the HW (from 2 s before HW to the onset of HW) was correlated with subsequent avoidance behavior, we calculated the log odds ratio from the trained classifier. The results showed that Type 2 neurons had a significantly higher log odds ratio, indicating a high correlation between unit activity and success in avoidance, whereas no discriminatory modulation was observed for Type 1 neurons or others (one sample t-test, t(298) = .564; p =.573, t(94) = 5.63; p <.001, t(238) = .531; p =.596, Type 1, Type 2, and others, respectively). These results indicate that Type 2 units carries information that is relevant for classifying trial as AW. To summarize, the populational neural activity of the mPFC in rats within the E-zone distinguished between two types of HW behavior, and increased activity of unit subgroups was strongly associated with the execution of AW.

Revised version: Based on the previous clustering analysis that identified functionally cohesive sub-populations of neurons (Types 1 and 2), we further investigated the predictive encoding by the ensemble activity of the two subpopulations. The naïve Bayesian classifier uses class likelihood, P(unit|AW) and P(unit|EW), to classify a particular type of withdrawal response. By comparing the class likelihoods, we can assess the ability of these subgroups to distinguish between AW and EW. This ability is quantified by a ’Class Discrimination Index,’ calculated as a log odds ratio. This index reflects the likelihood that a given feature set is associated with a particular type of withdrawal response (AW or EW). To test whether the unit activity right before the head-withdrawal (from 2 s before the head-withdrawal to the onset of the head-withdrawal) was correlated with subsequent avoidance behavior, we calculated the log odds ratio from the trained classifier. The results showed that Type 2 neurons had a significantly higher log odds ratio, indicating a high correlation between unit activity and success in avoidance, whereas no discriminatory modulation was observed for Type 1 neurons or others (one sample t-test, t(298) = .564; p =.573, t(94) = 5.63; p <.001, t(238) = .531; p =.596, Type 1, Type 2, and others, respectively). These results indicate that Type 2 units carries information that is relevant for classifying trial as AW. To summarize, the populational neural activity of the mPFC in rats within the E-zone distinguished between two types of head-withdrawal behavior, and increased activity of unit subgroups was strongly associated with the execution of AW.

Added in the revised version: Overlapping populations of mPFC neurons adaptively encode spatial information and defensive decision

Added in the revised version: Through a series of analyses, we found that both distance and foraging decision can be decoded from the same mPFC neural ensemble. PCA results support the that the mPFC ensemble activity forms a distinctive functional subspace when the rat is engaged in foraging decision facing the robotic predator in the E-zone. Previous studies have also observed different firing patterns from the same mPFC neural ensembles when the rat switches between tasks (Durstewitz et al., 2010; Nigro et al., 2023). Therefore, it would be interesting to see whether the distance and foraging decision are encoded by two exclusively subsets of neurons or by the whole ensemble.

Added in the revised version: To quantitatively assess each neuron’s contribution to the distance regressor and event classifier, we calculated the feature importance scores for each decoder (Figure 7A). This metric measures the increase in error when a neuron’s data is shuffled. For example, if a neuron holds substantial data and is heavily relied upon by the decoder, shuffling its signal will significantly increase the error. We analyzed the permutation feature importance scores for each neuron for both distance regression and event classification (see Materials and Methods).

Added in the revised version: Feature importance analysis found no evidence of dedicated neural subset for Distance/Event encoding.

Added in the revised version: (A) Schematic diagram showing computational protocols for the feature importance analysis

Added in the revised version: (B) Frequency distribution of all recorded units for their feature importance (measured in error increase) in distance encoding. Only a few units produced non-negligible increase in error when removed from the data indicating the absence of dedicated distance-encoding neurons. Red line indicates 95th percentile.

Added in the revised version: (C) Accuracy of distance regressor without top 20% of high-performance units. Even when 20% of units with high feature importance score were removed, the distance regressor could decode rat’s location.

Added in the revised version: (D) Frequency distribution of all recorded units for their feature importance (measured in accuracy drop) in event classification. Only a few units produced non-negligible decrease in accuracy when removed from the data indicating the absence of dedicated event-encoding neurons. Shuffling unit’s data resulted small decrease in error, and even some showed increase in accuracy. Red line indicates 95th percentile.

Added in the revised version: (E) Accuracy of event classifier without top 20% of high-performance units. Even when 20% of units with high feature importance score were removed, the event classifier could decode the type of rat’s defensive behavior.

Added in the revised version: (F) Correlation between distance regressor’s feature importance score and event classifier’s feature importance score. There was no correlation between feature importance of two types of decoding

Added in the revised version: First, for the distance regressor, the median permutation feature importance score was 0.73 cm (Figure 7B). This suggests that shuffling one unit’s neural activity increases the decoding error by less than 1 cm. Moreover, 95% of units exhibited a permutation feature importance score below 1.998 cm, indicating only a negligible increase. To determine whether a subset of neurons with high feature importance scores solely drives distance decoding, we tested whether removing the top 20% of neurons with the highest scores in a given session affects distance decoding. The results showed that although removal significantly increased the decoding error, the regressor trained with the remaining data still achieved a significantly lower error (17.35cm; Wilcoxon signed rank test, p<.001). Next, we applied a similar method to the event classifier. This analysis revealed a median permutation feature importance score of 0.01 and a 95th percentile of 0.06. We also tested the effect of removing the top 20% of neurons with the highest scores and found that the decoder still exhibited a significantly lower error (paired t-test, t(39)=6.899, p<.001). This result indicates that there is no sub-population which dominates in encoding distance information.

Added in the revised version: Furthermore, we tested whether there is a correlation between the feature importance scores of the distance regressor and the event classifier. If there were two subgroups of neurons, each dedicated to encoding distance and events respectively, we might expect to see a negative correlation. To verify this, we calculated Pearson’s correlation for each session. We found that there was no significant correlation between the feature importance scores of the two encoders (r=0.077, p=.062), indicating the absence of two distinct subpopulations that exclusively encode either distance or event.

Added in the revised version: Additionally, we found no evidence that Type I or Type II neurons have statistically higher or lower feature importance in the distance regressor. These results support the notion that distance and avoidance outcomes are not encoded separately by subsets of neurons; rather, the entire ensemble encodes two different concepts according to the given context.

Discussion

Earlier version: We found distinctively different states of ensemble activity from the mPFC (PL and IL) of a foraging rat, and these states were linked to specific areas or tasks within the foraging environment. Specifically, the rat’s distance could be decoded with stable accuracy during navigation, while its ensemble state switched to encode a defensive behavior when the rat entered the immediate vicinity of a reward and threat region (E-zone). Within the high-conflict area, the rat’s threat-evading behavior (AW), coupled with forgoing the reward, could be decoded using a naïve Bayesian classifier. These results demonstrated that the rat’s mPFC neurons were closely associated with two different functional modes of processing allocentric spatial information as well as egocentric information, the reward and threat.

Revised version: We found distinctively different states of ensemble activity from the mPFC (PL and IL) of a foraging rat, and these states were linked to specific areas or tasks within the foraging environment. Specifically, the rat’s distance from the conflict zone could be decoded with stable accuracy during navigation, while its ensemble state switched to encode a defensive behavior when the rat entered the immediate vicinity of a reward and threat region (E-zone). Within the high-conflict area, the rat’s threat-evading behavior (AW), coupled with forgoing the reward, could be decoded using a naïve Bayesian classifier. These results demonstrated that the rat’s mPFC neurons were closely associated with two different functional modes of processing allocentric spatial information as well as egocentric information, the reward and threat.

Earlier version: In the present study, the rat’s spatial information was extracted from the mPFC neuron in the form of distance. Previous studies examining the mPFC’s spatial representation found a spectrum of outcomes, from an absence of spatial information (Poucet, 1997) to accurate (x, y) coordinates with errors as small as a few centimeters (Ma et al., 2023; Mashhoori et al., 2018). In some studies, the spatial coding mechanism was directly compared between the hippocampus and the mPFC (Kaefer et al., 2020; Powell & Redish, 2014), and the mPFC cells were found to encode a generalized form of space rather than specific locations. Therefore, we built an ANN regressor targeting the distance, a more generalized spatial index which also reflects an internal representation of value and risk.

Revised version: In the present study, the rat’s spatial information was extracted from the mPFC neuron in the form of distance. Previous studies examining the mPFC’s spatial representation found a spectrum of outcomes, from an absence of spatial information (Poucet, 1997) to accurate (x, y) coordinates with errors as small as a few centimeters (Ma et al., 2023; Mashhoori et al., 2018). In some studies, the spatial coding mechanism was directly compared between the hippocampus and the mPFC (Kaefer et al., 2020; Powell & Redish, 2014), and the mPFC cells were found to encode a generalized form of space rather than specific locations. Therefore, we focused on distance as a more generalized spatial index that also reflects an internal representation of value and risk, and built an ANN regressor to target it.

Earlier version: The decoding analysis showed that brief navigation was sufficient to engage a majority of mPFC neurons recorded during foraging into spatial encoding. Furthermore, the spatial encoding within the F-zone showed a consistent representation, irrespective of navigational intent or visual cues during foraging (inbound vs. outbound). The spatial resolution encoded by the ensemble of neurons was precise (13.86 cm in the F-zone) considering the body length of the adult rat used in the experiment (20 cm or above excluding the tail). Therefore, we argue that mPFC neurons are consistent with the spatial encoding hypothesis.

Revised version: The decoding analysis showed that brief navigation was sufficient to engage a majority of mPFC neurons recorded during foraging into spatial encoding. Furthermore, the decoding accuracy was consistent within the F-zone, irrespective of navigational intent or visual cues during foraging (inbound vs. outbound). The spatial resolution encoded by the ensemble of neurons was precise enough to provide navigational utility (13.86 cm in the F-zone) considering the body length of an adult rat (∼20 cm excluding the tail). Therefore, we argue that the functional properties of mPFC neurons in the experiment are consistent with the spatial encoding hypothesis.

Earlier version: However, our observations also indicate that the spatial encoding in the mPFC was susceptible to disruption. Our distance regressor showed a higher error rate when rats engaged with the Lobsterbot inside the E-zone. A previous study also found that the presentation of a disfavored reward results in erroneous excursion of the location decoder near the feeder location (Mashhoori et al., 2018). Importantly, such disruption was not limited to the E-zone, appearing during non-navigational behaviors unrelated to reward or goal and significantly degrading the quality of spatial representation. These results suggest that spatial encoding in the mPFC was task-dependent.

Revised version: However, our observations also indicate that the spatial encoding in the mPFC was susceptible to disruption. Our distance regressor showed a higher error rate when rats engaged with the Lobsterbot inside the E-zone. A previous study also found that the presentation of a disfavored reward results in erroneous excursion of the location decoder near the feeder location (Mashhoori et al., 2018). Importantly, such disruption was not limited to the E-zone. Non-navigational behaviors such as grooming and sniffing greatly reduced the accuracy of location prediction suggesting that spatial encoding in the mPFC is tightly bound with task-demands such as goal-seeking and threat avoidance.

Earlier version: In the F-zone, where the regressor exhibited high spatial decoding accuracy, strategic navigation was primarily observed. Hok et al. (2005) argued that the mPFC neurons display spatial-selective firing patterns when the rat is engaged in navigation tasks rather than aimlessly wandering. Similarly, Ma et al. (2023) found that complex tasks involving working memory and reward improved spatial decoding accuracy in the mPFC. Consistent with these results, in our study, the removal of the Lobsterbot —which simplified the task to mere reward collection— resulted in decreased decoding accuracy. Given these observations, along with the mPFC’s lack of consistency in spatial encoding, it is plausible that the mPFC operates in multiple functional modes, and some of which involve spatial encoding.

Revised version: In the F-zone, where the regressor exhibited high spatial decoding accuracy, strategic navigation was primarily observed. Hok et al. (2005) argued that the mPFC neurons display spatial-selective firing patterns when the rat is engaged in navigation tasks rather than aimlessly wandering. Similarly, Ma et al. (2023) found that complex tasks involving working memory and reward improved spatial decoding accuracy in the mPFC. Consistent with these results, in our study, the removal of the Lobsterbot —which simplified the task to mere reward collection— resulted in decreased decoding accuracy. Given these observations, along with the mPFC’s lack of consistency in spatial encoding, it is plausible that the mPFC operates in multiple functional modes, and the spatial encoding mode is preempted when the complexity of the task requires deliberate spatial navigation.

Earlier version: Consistent with mPFC neurons being responsive to valence such as reward (Horst & Laubach, 2013), threat (Burgos-Robles et al., 2009), or both (Kim et al., 2018), the recorded neurons modulated their firing rate around the entry and withdrawal to the E-zone, where sucrose reward and robotic threat coexisted. As the mPFC is highly correlated with executive functions (Dalley et al., 2004; G. W. Diehl & Redish, 2023), we built an ensemble classifier and found that mPFC activity predicted the success of AW, possibly by integrating sensory information represented inside the region.

Revised version: Consistent with previous findings that mPFC neurons are sensitive to changes in valence such as reward (Horst & Laubach, 2013), threat (Burgos-Robles et al., 2009), or both (Kim et al., 2018), the recorded neurons modulated their firing rate around the entry and withdrawal to the E-zone, where sucrose reward and robotic threat coexisted. Consistent with the results indicating the mPFC is highly involved with executive functions (Dalley et al., 2004; G. W. Diehl & Redish, 2023), an ensemble classifier trained with the recorded unit activity predicted the success of AW at a fairly high accuracy (0.73). The prediction was significantly above the chance level (0.5), even several seconds before the action.

Added in the revised version: Interestingly, we found no difference between the PL and IL in the decoding accuracy of distance or avoidance decision. This somewhat surprising considering distinct roles of these regions in the long line of fear conditioning and extinction studies, where the PL has been linked to fear expression and the IL to fear extinction learning (Burgos-Robles et al., 2009; Dejean et al., 2016; Kim et al., 2013; Quirk et al., 2006; Sierra-Mercado et al., 2011; Vidal-Gonzalez et al., 2006). On the other hand, more Type 2 neurons were found in the PL and more Type 1 neurons were found in the IL. To recap, typical Type 1 neurons increased the activity briefly after the head entry and then remained inhibited, while Type 2 neurons showed a burst of activity during head entry and sustained increased activity. One study employing context-dependent fear discrimination task (Kim et al., 2013) also identified two distinct types of PL units: short-latency CS-responsive units, which increased firing during the initial 150 ms of tone presentation, and persistently firing units, which maintained firing for up to 30 seconds. Given the temporal dynamics of Type 2 neurons, it is possible that our unsupervised clustering method may have merged the two types of neurons found in Kim et al.’s study.

Added in the revised version: While we did not observe decreased IL activity during dynamic foraging, prior studies have shown that IL excitability decreases after fear conditioning (Santini et al., 2008), and increased IL activity is necessary for fear extinction learning. In our paradigm, extinction learning was unlikely, as the threat persisted throughout the experiment. Future studies with direct manipulation of these subpopulations, particularly examining head withdrawal timing after such interventions, could provide insight into how these subpopulations guide behavior.

Earlier version: We also found that the high activity levels of Type 2 neurons were highly correlated with the success of AW. What is the role of the Type 2 sub-population in the execution of active defensive behavior? One straightforward hypothesis is that the increased activity simply represents fear or anxiety motivation. This premise is supported by multiple experiments where the PL is vital for the expression of freezing responses (Burgos-Robles et al., 2009; Kim et al., 2018). Consequently, Type 2 units may represent the perceived threat level, which would strongly correlate with the rat’s voluntary HW. Alternatively, the Type 2 sub-population may contribute to the computation or decision-making processes underlying defensive behaviors. This is based on a study by Halladay and Blair (2015), where they identified two sub-populations selectively responding to specific defensive behaviors, namely, conditioned movement inhibition and excitation. Because Type 2 neurons increased their activity in AW trials, it is possible that they overlap with the excitation-related neuronal subgroups in Halladay and Blair’s study.

Revised version: We also found that the activity of Type 2 neurons was highly correlated with the success of AW. What is the role of the Type 2 sub-population in the execution of active defensive behavior? One straightforward hypothesis is that the increased activity simply represents internal motivation of fear or anxiety. This premise is supported by multiple experiments where the PL is vital for the expression of freezing responses (Burgos-Robles et al., 2009; Kim et al., 2018). Consequently, Type 2 units may represent the perceived threat level, which would strongly correlate with the rat’s voluntary head-withdrawal. Alternatively, the Type 2 sub-population may contribute to the computation or decision-making processes underlying defensive behaviors. This is based on a study by Halladay and Blair (2015), where they identified two sub-populations selectively responding to specific defensive behaviors, namely, conditioned movement inhibition and excitation. Because Type 2 neurons increased their activity in AW trials, it is possible that they overlap with the excitation-related neuronal subgroups in Halladay and Blair’s study. Despite this interesting conjecture, any analysis based on recording data is only correlational, mandating further studies with direct manipulation of the subpopulation to confirm its functional specificity.

Encoded representations in the mPFC range from prediction of upcoming decisions (Passecker et al., 2019) to rule switching (Laskowski et al., 2016). This multiplicity of cognitive functions in the mPFC has also been observed in primates, giving rise to the adaptive coding model (Duncan, 2001). This model posits that the mPFC neurons are fine-tuned to encode task-relevant concepts within specific contexts to meet the requirements of the current task. In our study, more than 80% of units—which coded distance information in the F-zone—modulated with the defensive decision making when entering or leaving the E-zone, thus supporting the adaptive coding model and functional switching between different contexts. The PCA results further revealed the heterogeneity of the populational activity in the E-zone as its manifold, or subspace of possible neural activity (for review, see Ebitz & Hayden, 2021), was distinctive from those of the F- and N-zones.

Earlier version: It is worth emphasizing that upon exiting the E-zone the neurons resumed encoding distance in the same manner as in previous trials. This is substantiated by the fact that our regressor was trained on whole session data rather than trial-by-trial data, and it would have been impossible to accurately decode distance if the neurons changed their encoding method every time the rats re-entered the F-zone. Interestingly, in our experiment paradigm, the rats quickly shifted between contexts—navigational in the F-zone vs. confrontational in the E-zone—within the same session. The consistent spatial decoding in the F-zone throughout the session indicates that mPFC neurons selectively switch functions, rather than being randomly activated when the context changes.

Revised version: It is worth emphasizing that upon exiting the E-zone the neurons resumed encoding distance in the same manner as they did in previous trials. This is substantiated by the fact that our regressor was trained by the recording data from the whole session and therefore, it would have been impossible to accurately decode the distance if the neurons changed their encoding every time the rats re-entered the F-zone. Interestingly, in our foraging task, the rats quickly shifted their behavior mode between contexts—navigational in the F-zone vs. confrontational in the E-zone—within the same session. The temporal stability of spatial decoding in the F-zone throughout the session, despite the intermittent engagement in confrontational behavior in the E-zone, indicates that mPFC neurons selectively switch functions, rather than being randomly activated by the context change.

Earlier version: In summary, our findings reveal a dual, context-sensitive functionality in the rat’s mPFC, accommodating both spatial navigation and defensive decision-making during dynamic foraging. We demonstrated that the mPFC co-represented spatial and foraging-related information—such as reward, threat, and execution of defensive behavior—depending on the context (Figure 7). This functional switch not only highlights the mPFC’s multifaceted encoding capabilities but also reconciles conflicting models regarding its role in spatial orientation. Future studies should be performed to determine whether these functional shifts are governed by top-down processes involving specific neural arbitrators, or by bottom-up processes driven by signal strength.

Revised version: In summary, our findings reveal a dual, context-sensitive encoding in the rat’s mPFC, accommodating both spatial navigation and defensive decision-making during dynamic foraging. We demonstrated that the mPFC co-represented spatial and active foraging-related information depending on the context (Figure 8A) Figure. This functional switch not only highlights the mPFC’s multifaceted encoding capabilities but also reconciles conflicting models regarding its role in spatial orientation.

Added in the revised version: Hypothetical control models by which mPFC neurons assume different functional states.

Added in the revised version: (A) In the F-zone, navigational behaviors enhance the mPFC’s encoding of spatial information compared to other zones. In the N-zone, spatial coding diminishes when the rat engages in non-navigational behaviors. However, in the E-zone, these neurons shift their encoding strategy and become involved in coding for active foraging. We did not find a subset of neurons dedicated exclusively to either spatial coding or active foraging throughout the session. Instead, neurons changed their encoding scheme in a population-wide manner.

Added in the revised version: (B) Two hypotheses about how the switch is manifested. In this example, most mPFC neurons encode spatial information (blue circles). Information encoded in the mPFC can be regulated by internal/external arbitration signal (top-bottom blue arrow from green circles), or influenced by direct sensory inputs and navigation-related signals (left-right blue arrow) that prompt mPFC neurons to encode spatial information.

Added in the revised version: In primates and rodents, the mPFC is often likened to a “department store” due to its involvement in numerous cognitive functions (e.g., reward, decision-making, prediction error). One pioneering theory suggests that the mPFC’s involvement in diverse functions stems from its ability to learn associations between context, locations, events, and adaptive responses (Euston et al., 2012). However, in our multi-task design, decoder accuracy varied with context and task demands, prompting the question: how does the mPFC decide which information or associations to encode?

Added in the revised version: Two main hypotheses could explain this switch. A bottom-up hypothesis suggests sensory inputs or upstream signals dictate encoding priorities, while a top-down hypothesis proposes that an internal or external “arbitrator” selects the encoding mode and coordinates the relevant information (Figure 8B). Although the current study is only a first step toward finding the regulatory mechanism behind this switch, our control experiment, where rats reverted to a simple shuttling task, provide evidence that might favor the top-down hypothesis. The absence of the Lobsterbot degraded spatial encoding rather than enhancing it, indicating that simply reducing the task demand is not sufficient to activate one particular type of encoding mode over another. The arbitrator hypothesis asserts that the mPFC neurons are called on to encode heterogenous information when the task demand is high and requires behavioral coordination beyond automatic, stimulus-driven execution. Future studies incorporating multiple simultaneous tasks and carefully controlling contextual variables could help determine whether these functional shifts are governed by top-down processes involving specific neural arbitrators or by bottom-up signals.

Removed in the revised version: mPFC neurons switch between different functional modes depending on the location.

Removed in the revised version: In the F-zone, goal-directed navigation prompts a greater number of mPFC neurons to encode spatial information compared to other zones. In the E-zone, these neurons switch the encoding scheme for reward- or threat-related information.

Materials and Methods

Subjects

Eight male Sprague Dawley rats (250-300 g, 5-6 weeks of age, Orient Bio.) were used in this study. The rats were housed in pairs on a 12-h light/dark cycle with food and water available ad libitum for a one-week acclimation period. In the following week, rats were handled for 10 min each day until they displayed reduced fear responses towards the experimenter.

Apparatus and procedures

The apparatus followed the same dimensions as in Kimm & Choi (2018), measuring 800 mm × 480 mm and a height of 500 mm excluding the robot compartment. It was covered in gray EMF-shield fabric to reduce high-frequency noise during experiments. Battery-powered LED strip lighting illuminated the apparatus and −60 dB white noise masked background sounds. The system employed four Arduino Uno boards and a Raspberry Pi board for experiment control and behavioral tracking. All sensor and control signals were relayed to an RV5 multiprocessor (Tucker-Davis Technologies, FL, USA) for further analysis. Bird’s-eye view videos were captured during sessions, and the rat’s location was tracked using the butter package (https://github.com/knowblesse/butter). All control schematics and codes are publicly available (https://github.com/knowblesse/Lobster/Apparatus).

Earlier version: Before the training, food and water were gradually restricted until the rats’ weight reached 80% of their initial value. Over three days, they acclimated to the apparatus, daily confined from the robot compartment for 20 min. Training began with the H phase, during which rats had unrestricted sucrose port access for 20 min. Upon reaching 1,500 licks per trial, rats proceeded to the three-day S phase. The S phase was 30 min long, and the gate was closed 6 s after the first lick, forcing rats into the N-zone to reopen it. The subsequent L phase introduced the Lobsterbot, which executed two grabbing motions after 3 s (30%) or 6 s (70%) from the initial lick. Lasting three days, rats attempting fewer than 30 approaches in 30 min were excluded from further experiments.

Revised version: Before the training, food and water were gradually restricted until the rats’ weight reached 80% of their initial value. Over three days, they acclimated to the apparatus for 20 min. Training began with the habituation phase, during which rats had unrestricted sucrose port access for 20 min. The sucrose port has an IR sensor which was activated by a single lick. The rat usually stays in front of the lick port and continuously lick up to a rate of 6.3 times per second to obtain sucrose. Any sucrose droplets dropped in the bottom sink were immediately removed by negative pressure so that the rat’s behavior was focused on the licking. Upon reaching 1,500 licks per trial, rats proceeded to the three-day shuttling phase. The shuttling session was 30 min long, and the gate was closed 6 s after the first lick, forcing rats into the N-zone to reopen it. The subsequent Lobsterbot phase introduced the Lobsterbot, which executed two grabbing motions after 3 s (30%) or 6 s (70%) from the initial lick. After three days of Lobsterbot sessions, the rats underwent microdrive implant surgery, and recording data were collected from subsequent sessions, either Lobsterbot or shuttling sessions, depending on the experiment. For all post-surgery sessions, those with fewer than 20 approaches in 30 minutes were excluded from further analysis.

Surgery

Earlier version: Before surgery, rats were anesthetized using isoflurane (3%-5%) and oxygen (500 mL/min). Enrofloxacin (5 mg/kg) and meloxicam (1 mg/kg) were administered to prevent infection and mitigate postsurgical pain. A craniotomy was performed at AP +3.2 mm and ML 0.5 mm from the bregma using a carbide burr (Komet, Kempten, Germany) mounted on a high-speed drill (Strong 207A; Saeshin, Daegu, Korea). The tip of each microdrive was lowered to the following coordinates: AP +3.2 mm, ML right 0.3 ∼ 0.6 mm, DV -4 mm from the bregma. The precise ML coordinate was determined after the craniotomy to avoid damage to the superior sagittal sinus. Six to seven screws were securely affixed to the skull, and two additional screws above the cerebellum were used as the ground and reference points. Finally, dental cement was applied around the microdrive to secure the implant. Rats were allowed at least one week of recovery before the recordings. All procedures were approved by the Korea University Institutional Animal Care & Use Committee (KUIACUC-2020-0029).

Revised version: Before surgery, rats were anesthetized using isoflurane (3%-5%) and oxygen (500 mL/min). Enrofloxacin (5 mg/kg) and meloxicam (1 mg/kg) were administered to prevent infection and mitigate postsurgical pain. A craniotomy was performed at AP +3.2 mm and ML 0.5 mm from the bregma using a carbide burr (Komet, Kempten, Germany) mounted on a high-speed drill (Strong 207A; Saeshin, Daegu, Korea). The tip of each microdrive was lowered to the following coordinates: AP +3.2 mm, ML right 0.3 ∼ 0.6 mm, DV −4 mm from the bregma. The precise ML coordinate was determined after the craniotomy to avoid damage to the superior sagittal sinus. Six to seven screws were securely affixed to the skull, and two additional screws above the cerebellum were used as the ground and reference points. Finally, dental cement was applied around the microdrive to secure the implant. Rats were allowed at least one week of recovery before the recordings. All procedures were approved by the Korea University Institutional Animal Care & Use Committee.

Data collection

Earlier version: Two types of microdrives were used in the recordings: a custom-made 4-tetrode microdrive and a pre-assembled 16-channel silicon probe (A2x2-tet-10 mm-150-150-121-HZ16_21 mm, NeuroNexus, MI, USA) mounted on a 3D printed microdrive. The microdrive featured an M1.6-18 mm screw with a 0.35 mm pitch. The recording site was retroactively determined by calculating the number of rotations from the final recording site, identified by a marking lesion.

Revised version: Two types of microdrives were used in the recordings: a custom-made 4-tetrode microdrive and a pre-assembled 16-channel silicon probe (A2×2-tet-10 mm-150-150-121-HZ16_21 mm, NeuroNexus, MI, USA) mounted on a 3D printed microdrive. The microdrive featured an M1.6-18 mm screw with a 0.35 mm pitch. The recording site was retroactively determined by calculating the number of rotations from the final recording site, identified by a marking lesion.

During each recording session, the rat’s microdrive was connected to a ZIF-Clip digital headstage (Tucker-Davis Technologies, FL, USA). The neural signals were collected at a 25 kHz sampling rate, amplified, digitized from the head stage, and transferred to the RZ5 (Tucker-Davis Technologies, FL, USA) for data acquisition. The acquired data stream was further bandpass filtered between 300 and 5000 kHz, and all spikes above 3 to 4 sigma compared to the stream were collected for further clustering analysis. The Offline Sorter (Plexon, TX, USA) was used for the manual clustering process. All units with firing rates below 0.5 Hz or exactly 60 Hz (power noise) were considered artifacts and removed from further analysis. Sessions with fewer than 10 recorded units were also excluded.

Spatial decoding data preparation

For the neural data used in the spatial encoding analysis of the mPFC, the following method was employed. First, a Gaussian kernel with a standard deviation of 100 ms and a length of 1000 ms was prepared. Each unit’s spike data were down-sampled and serialized to a 1000 Hz vector, and the kernel was applied to smooth the signal. The mean and standard deviation values of this vector were computed and used to normalize the vector. After normalization, the vector was binned into 50 ms windows, resulting in 20 data points per second. Finally, each datapoint was paired with the corresponding location data.

Added in the revised version: The spatial firing map was constructed using the same normalized activity data. Neural activity was averaged for each pixel and then smoothed using a Gaussian filter with a sigma of 15 px and a filter size of 1001 px.

ANN regressor

A neural network regressor was built using the PyTorch package. It comprised seven layers in total, including the output layer; four of these were fully connected (FC) layers and two were drop-out layers. The initial weights for the FC layers were normalized with a mean of zero and a standard deviation of 0.2. Mean squared error loss was used in conjunction with the stochastic gradient descent method. For each regressor, a control regressor was built, which used the same data-label set, but the relationships were randomly shuffled. The evaluation of the regression results used a stratified 5-fold cross-validation while ensuring that each subset maintained the same proportion of data from the three zones. The five regression results per dataset were combined, and these results were compared with the true dataset using MAE. For the training/testing dataset, preprocessed 50 ms-long neural data were used. The binned neural data from all units (n) were concatenated to form a 1 × n size vector. These vectors were paired with the corresponding distance during the 50 ms time window. In the dataset adjusted for uneven location visits, we divided distance values into five equally sized bins. Then, a sub-dataset was created that contains an equal number of data points for each of these bins. Two different GPU servers were employed for training: a Tesla V100 from the National IT Industry Promotion Agency (NIPA, Seoul, Korea) and an in-house A4000.

Decoding error heatmap construction

To construct the decoding error heatmap, the following method was used. First, error data from each pixel were collected and mean values were computed. To accommodate pixels not visited by a rat during an entire session, an interpolation method was employed. A custom Matlab script, built using the scatteredInterpolant function, was used to interpolate errors for unknown locations. The 2D interpolated decoding error was subsequently smoothed using a Gaussian filter with a sigma value of 15 px and a filter size of 1001 px. These smoothed error values were then converted to cm, and a contour plot was overlaid using 25 levels.

PCA

Earlier version: A custom Python script for PCA was created using the sci-kit-learn package. Neural data vectors generated for the spatial decoding analysis were used for this analysis. The behavior and head tracking data were matched with neural data vectors, and the corresponding zone was labeled to the vectors. All neural data vectors from a single session were projected onto a new space using the PCA. For each zone, the centroid of the projected neural data vectors was calculated using Euclidean distance. The within-zone distance and the between-zone distance were calculated using the equations below. where Cz denotes the centroid of the zone z and xzrepresents a neural data vector from the zone z.

Revised version: A custom Python script for PCA was created using the sci-kit-learn package. Neural data vectors generated for the spatial decoding analysis were used for this analysis. The behavior and head tracking data were matched with neural data vectors, and the corresponding zone was labeled to the vectors. All neural data vectors from a single session were projected onto a new space using the PCA. For each zone, the centroid of the projected neural data vectors was calculated using Euclidean distance. The within-zone distance and the between-zone distance were calculated using the equations below. where CZ denotes the centroid of the zone z and xZ represents a neural data vector from the zone z.

Representative figures were drawn in a manner to highlight the cluster differences. Among all projected neural vectors designated to a zone, the 2000 vectors with the shortest distance from the centroid were selected. From these vectors, 200 were randomly chosen and the first two dimensions of the projected value were plotted.

Behavior-responsive unit analysis

Earlier version: Behavior-responsive units were identified using the following method. For each trial, spikes within a 4-s window (2 s before and after the onset of three events: gate opening upon returning to the N-zone, HE, and HW) were counted and registered in 50 ms-long time bins, resulting in 80 data points per trial. These neural activities were then averaged across all trials, producing a 1 × 80 vector for each event. The binned neural activity ± 2 s from the entrance to the N-zone was chosen as a baseline, considering that the rat was consistently oriented towards areas other than the robot. The mean and standard deviation of the baseline were used to normalize activity during HE and HW. The overall method created two event vectors, HE-vector and HW-vector, per unit. A unit was considered behavior-responsive to an event if any value within the 80 data points of the event vector was greater or less than 3 sigma.

Revised version: Behavior-responsive units were identified using the following method. For each trial, spikes within a 4-s window (2 s before and after the onset of three events: gate opening upon returning to the N-zone, head-entry, and head-withdrawal) were counted and registered in 50 ms-long time bins, resulting in 80 data points per trial. These neural activities were then averaged across all trials, producing a 1 × 80 vector for each event. The binned neural activity ± 2 s from the entrance to the N-zone was chosen as a baseline, considering that the rat was consistently oriented towards areas other than the robot. The mean and standard deviation of the baseline were used to normalize activity during the head-entry and the head-withdrawal. The overall method created two event vectors, head-entry-vector and head-withdrawal-vector, per unit. A unit was considered behavior-responsive to an event if any value within the 80 data points of the event vector was greater or less than 3 sigma.

Hierarchical cell clustering analysis

Earlier version: To cluster cells that exhibited similar responses during the HE and/or HW, an unsupervised hierarchical cell clustering method was employed. For every pair of units, the similarity score using Pearson’s r was calculated, and the pair with the highest similarity score was grouped to form a new pseudo-unit that had the average response of the two units. This procedure was repeated until all units were grouped into eight pseudo-units. Among each pseudo-unit, if the number of units belonging to the group was equal to or greater than 50, the pseudo-unit was considered a valid cluster. Otherwise, the pseudo-unit was classified as an artifact. Although this method had a chance of losing small population groups that exhibited important unique activity during the event, other groups were focused on as they consistently appeared across multiple recording sessions and subjects.

Revised version: To cluster cells that exhibited similar responses during the head-entry and/or head-withdrawal, an unsupervised hierarchical cell clustering method was employed. For every pair of units, the similarity score using Pearson’s r was calculated, and the pair with the highest similarity score was grouped to form a new pseudo-unit that had the average response of the two units. This procedure was repeated until all units were grouped into eight pseudo-units. Among each pseudo-unit, if the number of units belonging to the group was equal to or greater than 50, the pseudo-unit was considered a valid cluster. Otherwise, the pseudo-unit was classified as an artifact.

Added in the revised version: To compensate for the limited number of neurons recorded per session, the hyperparameter set was chosen to generalize their activity and categorize them into major types, allowing us to focus on neurons that appeared across multiple recording sessions. Although changes in the hyperparameter sets resulted in different numbers of clusters, the major activity types remained consistent (Supplementary Figure S8). However, there is a chance that this method may not differentiate smaller subsets of neurons, particularly those with fewer than 50 recorded neurons.

Run-and-stop event analysis

To examine whether a change in movement velocity was related to neural activity, a “run-and-stop” event was defined by the following criteria: 1) to qualify as “run” at a given moment, the rat must be moving at least 7.5 cm/s for at least 2 s or longer; 2) to qualify as the “stop” at a given moment, the rat must be moving at most 5 cm/s for at least 1 s or longer; and 3) the “run-and-stop” event was the time point between “run” and “stop” outside the E-zone. The temporal relationship between the run-and-stop event and neural activity was visualized using the same method as in the behavior-responsive unit analysis.

Earlier version: Naïve Bayes HW-type classifier and data preparation

Revised version: Naïve Bayes AW/EW classifier and data preparation

Earlier version: For the neural data used in HW-type classification analysis, a data preparation method similar to that of the spatial decoding analysis was employed. First, the same Gaussian kernel was generated and applied to the down-sampled and serialized neural data. Mean and standard deviation values of this vector were computed for later normalization. Unlike the 50 ms-length neural data used in spatial regression, 2-s-length neural data was used. For each trial, all spikes within a 2-s window around the onset of HW were collected and serialized into a 1000 Hz vector. The Gaussian kernel was applied and normalized using precalculated values. This vector was then binned into 50 ms intervals, resulting in 40 data points per event. Binned vectors from all recorded units (n) were concatenated to form a 1 × (40n) size vector per trial. These vectors were then paired with their corresponding HW-type, AW or EW. In the analysis using varying time windows, the previously employed data preparation method was the same, but the time range of the neural data varied relative to the HW.

Revised version: For the neural data used in the head-withdrawal-type (AW/EW) classification analysis, a data preparation method similar to that of the spatial decoding analysis was employed. First, the same Gaussian kernel was generated and applied to the down-sampled and serialized neural data. Mean and standard deviation values of this vector were computed for later normalization. Unlike the 50 ms-length neural data used in spatial regression, 2-s-length neural data was used. For each trial, all spikes within a 2-s window around the onset of the head-withdrawal were collected and serialized into a 1000 Hz vector. The Gaussian kernel was applied and normalized using precalculated values. This vector was then binned into 50 ms intervals, resulting in 40 data points per event. Binned vectors from all recorded units (n) were concatenated to form a 1 × (40n) size vector per trial. These vectors were then paired with their corresponding head-withdrawal-type, AW or EW. In the analysis using varying time windows, the previously employed data preparation method was the same, but the time range of the neural data varied relative to the head-withdrawal.

Earlier version: A naïve Bayesian classifier based on the multivariate Bernoulli model was built using the sci-kit-learn package. A custom Python script was written to train and evaluate the classifiers. The classifier employed a uniform prior and smoothing parameter, alpha, which was set to 1 in all classifiers used in the study. Similar to the analysis using the ANN regressor, a control classifier was built that was trained with shuffled data. The evaluation of the classifier results employed the stratified 5-fold cross-validation method with each subset maintaining the same proportion of HE/HW. Five classification results per dataset were combined, and these results were compared with the true dataset using balanced accuracy, calculated using the equation below. To interpret how the classifier anticipated upcoming HW behavior, class discrimination index was employed, focusing on a dataset with a time window of -2 to 0 s. To calculate the index, the class likelihood for every dimension was extracted during each test dataset evaluation, and a natural logarithm was applied to balance the small likelihood values. With each unit’s neural data divided into 40 bins, the log likelihood values were added to create a joint probability. This process yielded five joint likelihood values per unit, and their average was computed. Following Mladenić & Grobelnik’s (1999) method, the log odds ratio between P(unit|AW) and P(unit|EW) was calculated for all units using the equation below. where P(unit|AW) denotes joint likelihood (probability of observing unit activity given the AW) and P(unit|EW) denotes the same for the EW. These values were then compared according to the unit’s assigned group: Type 1, Type 2, or others. Taking the natural log of the odds ratio brings about a symmetricity in the results. A positive value represented a unit’s relatively increased activity during the AW compared to the EW, and a negative value represented the opposite.

Revised version: A naïve Bayesian classifier based on the multivariate Bernoulli model was built using the sci-kit-learn package. A custom Python script was written to train and evaluate the classifiers. The classifier employed a uniform prior and smoothing parameter, alpha, which was set to 1 in all classifiers used in the study. Similar to the analysis using the ANN regressor, a control classifier was built that was trained with shuffled data. The evaluation of the classifier results employed the stratified 5-fold cross-validation method with each subset maintaining the same proportion of head-entry/head-withdrawal. Five classification results per dataset were combined, and these results were compared with the true dataset using balanced accuracy, calculated using the equation below. To interpret how the classifier anticipated upcoming head-withdrawal behavior, class discrimination index was employed, focusing on a dataset with a time window of −2 to 0 s. To calculate the index, the class likelihood for every dimension was extracted during each test dataset evaluation, and a natural logarithm was applied to balance the small likelihood values. With each unit’s neural data divided into 40 bins, the log likelihood values were added to create a joint probability. This process yielded five joint likelihood values per unit, and their average was computed. Following Mladenić & Grobelnik’s (1999) method, the log odds ratio between P(unit|AW) and P(unit|EW) was calculated for all units using the equation below. where P(unit|AW) denotes joint likelihood (probability of observing unit activity given the AW) and P(unit|EW) denotes the same for the EW. These values were then compared according to the unit’s assigned group: Type 1, Type 2, or others. Taking the natural log of the odds ratio brings about a symmetricity in the results. A positive value represented a unit’s relatively increased activity during the AW compared to the EW, and a negative value represented the opposite.

Statistical analysis

Earlier version: Statistical analyses were performed using Prism (GraphPad, MA, USA) software. All paired t-tests used two-tailed p-values to assess the significance of differences. A p-value of <.05 was considered statistically significant. Both ANOVA and mixed-effect model analysis applied the Geisser-Greenhouse correction to account for violations of the sphericity assumption. The corrected degree of freedom is presented alongside F values. For multiple comparison tests, Sidak’s correction was used to control the family-wise error rate. For multiple linear regression analysis of the distance decoding error, the following model was incorporated. The model fitting method incorporated the least squares method assuming a Gaussian distribution of residuals. Boolean variables, presence of Lobsterbot, and recording site used dummy codes (0: Lobster, 1: Control; 0: PL, 1: IL) to represent the state of the experiment. All error bars in graphs represent one standard error of the mean. Statistically significant differences are denoted as * for p <.05, ** for p <.01, and *** for p <.001.

Revised version: Statistical analyses were performed using Prism (GraphPad, MA, USA). To confirm normality of the data, Shapiro-Wilk test was applied to all data we obtained from the experiment. For all statistical tests, p-value of <.05 were used to assess the significance of differences and all t-tests were two tailed. The correlation analysis comparing feature importance scores from two decoders excluded the 1% of outliers before computing Pearson’s r. Both ANOVA and mixed-effect model analysis applied the Geisser-Greenhouse correction to account for violations of the sphericity assumption. The corrected degree of freedom is presented alongside F values. For multiple comparison tests, Sidak’s correction was used to control the family-wise error rate. For multiple linear regression analysis of the distance decoding error, the following model was incorporated. The model fitting method incorporated the least squares method assuming a Gaussian distribution of residuals. Boolean variables, presence of Lobsterbot, and recording site used dummy codes (0: Lobster, 1: Control; 0: PL, 1: IL) to represent the state of the experiment. All error bars in graphs represent one standard error of the mean. Statistically significant differences are denoted as * for p <.05, ** for p <.01, and *** for p <.001.